# Programming Logic and Design Laboratory (LBYCPA1)

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or `YOUR ANSWER HERE`, as well as your name and other information below inside the double quotes. Remove the line `raise NotImplementedError()` once you have answered the problem.

REMINDERS: Section 4.15 of the Student Handbook (SH) 2018-2021 provides that:

As a Catholic University that prepares its students to become professionals and be the resource for Church and nation, De La Salle University puts a premium on academic honesty. A student's academic requirements, such as assignments, term papers, computer programs/projects, and thesis papers, should be his/her own work. He/She must distinguish between his/her own ideas and those of other authors.

The student must cite references, direct quotes, and other sources (including data obtained from tables, illustrations, figures, pictures, images, and video) following the prescribed format of the discipline. Should a student intend to submit a similar previous work for credit, in different courses, a written permission must first be obtained from the faculty members concerned. In addition, for works where he/she is only one of the authors, the following must be observed.

4.15.1 Provide a disclaimer stating that the work is not entirely his/hers;

4.15.2 The work cannot be published without written consent of the group of authors.

4.15.3 If the work is separable (i.e. parts are clearly attributed to the individual authors), the student may use and must cite his/her own part only; otherwise, the student must seek the written consent of the group of authors.

4.15.4 Computer codes obtained from open source reference may be copied freely, but the source should be cited, provided further that the faculty member allows the use of open source reference in the course.

## Rubric for the Final Laboratory Report

|Criteria|EXEMPLARY (90-100)|SATISFACTORY (80-89)|DEVELOPING (70-79)|BEGINNING (Below 70)|WEIGHT|
|:-|:-|:-|:-|:-|:-:|
|**Experimental Plan (Flowchart/ Algorithm)**|Experimental plan has supporting details and diagram/algorithm that is stated and well explained|Experimental plan has supporting details and diagram/algorithm that is stated but not explained|Experimental plan is vague or brief. It has supporting details and doesn’t have diagram/algorithm|No experimental plan presented|30%|
|**Conducting Experiment/ Activity**|Objective and Materials used (modern engineering tools, software, and instruments/equipment) are identified. Steps are easy to follow for conducting and experiment/activity|Objective is not stated. Materials used (modern engineering tools, software, and instruments/equipment) are identified. Steps are easy to follow for conducting and experiment/activity|Does not provide enough information to conduct an experiment/activity|No Objective, Materials used (modern engineering tools, software, and instruments/equipment), and steps for conducting experiment/activity provided|20%|
|**Codes/Data/ Program**|Data is well utilized in the program. Program codes are easy to read. Program output has no error. Questions are answered completely and correctly|Data is somewhat utilized in the program. Program code are easy to read. Program output has an output but logically incorrect. Some questions are answered completely and correctly|Data is not utilized in the program. It has a missing significant code/syntax in the program|No program presented|30%|
|**Grammar, logical presentation, and format**|The report was grammatically correct, logically presented and used the required format|The report had minimal grammatical errors and somewhat presented logically. The required format was used|The report had a lot of grammatical errors and not logically presented; the required format was barely used|The report had a lot of grammatical errors, was not logically presented and the required format was not used|20%|

```python
# Fill-up the following information (put inside double quotes)

NAME = "DELA CRUZ, JOHN CARLO THEO S."                   # e.g. LIGUTAN, DINO DOMINIC F.
STUDENT_NUMBER = "12195758"         # e.g. 20192118
SECTION = "EQ1"                # e.g. EQ1
DATE_TIME_COMPLETED = "06/07/22"    # e.g. 03/21/2022,11:00
```

<hr style="border:1px solid black"> </hr>

# Problem Set 8
<hr style="border:1px solid gray"> </hr>

1. Given a string which contains a long word, write a function that returns a dictionary of the letters used as keys (capitalized letters) with values counting the occurrences of that letter in the given long word. If the input string contains two or more words, the function should only consider the longest word. Note: Consider uppercase and lowercase letters to be the same.

<!-- grade_id: cell-32216f35a45b784e, points: 0, grade: false, solution: true -->
```python
def wordStatistics(sentence):
    if type(sentence) is not str:
        raise TypeError("Input not a valid string")
        
    longestWord = ''
    result = {}
    sentence = sentence.rstrip('.')
    sentence = sentence.upper()
    word = sentence.split(' ')
    for i in sentence:
        if len(i) > len(longestWord):
            longestWord = i
    for letters in longestWord:
        result[letters] = longestWord.count(letters)
    return result
```

<!-- grade_id: cell-a863d6410ce95310, points: 1, grade: true, solution: false -->
```python
# DO NOT MODIFY, THIS IS UNIT TESTING! [1 pt]
from nose.tools import assert_equal

assert_equal(wordStatistics("The ethics of xenotransplantation are relatively unworrying."), 
             {'X': 1, 'N': 4, 'O': 2, 'I': 1, 'P': 1, 'S': 1, 'R': 1, 'E': 1, 'L': 1, 'A': 3, 'T': 3})
assert_equal(wordStatistics("Her pulchritudinous likeness adorns everything from bestselling books to ski boards."), 
             {'N': 1, 'I': 2, 'O': 1, 'P': 1, 'H': 1, 'D': 1, 'S': 1, 'R': 1, 'C': 1, 'T': 1, 'L': 1, 'U': 3})
assert_equal(wordStatistics("Floccinaucinihilipilification"), 
             {'N': 3, 'I': 9, 'O': 2, 'P': 1, 'H': 1, 'F': 2, 'C': 4, 'T': 1, 'L': 3, 'A': 2, 'U': 1})
assert_equal(wordStatistics("Pseudopseudohypoparathyroidism needs no treatment, but genetic counselling is recommended."), 
             {'O': 4, 'I': 2, 'P': 4, 'D': 3, 'H': 2, 'M': 1, 'S': 3, 'R': 2, 'E': 2, 'Y': 2, 'T': 1, 'A': 2, 'U': 2})
assert_equal(wordStatistics("Pneumonoultramicroscopicsilicovolcanoconiosis"), 
             {'N': 4, 'O': 9, 'I': 6, 'P': 2, 'V': 1, 'M': 2, 'S': 4, 'R': 2, 'E': 1, 'C': 6, 'T': 1, 'L': 3, 'A': 2, 'U': 2})
```

<!-- grade_id: cell-b25b3cbf30198389, points: 1, grade: true, solution: false -->
```python
# HIDDEN TEST [1 pt]
```

2. A palindrome is a word, number, phrase, or other sequence of characters which reads the same backward as forward such as the number 256652. Write a Python function that accepts a list of either an integer or string sequence and returns a dictionary whose keys are the input sequences of the list and the values are tuple consisting of two possible palindromes that can be made from the input sequence.

<!-- grade_id: cell-2a67d9bb9ce95749, points: 0, grade: false, solution: true -->
```python
def toPalindromes(seq_list):
    if type(seq_list) is not list:
        raise TypeError("Input not a valid list")
    
    result = {}
    for i in seq_list:
        if type(i) == int:
            var = str(i)
            result[int(var)]=(int(var+"{}".format(var[::-1])), int("{}".format(var[::-1])+var))
        else:
            var = str(i)
            result[var] = (var+"{}".format(var[::-1]), "{}".format(var[::-1])+var)
    return result

```

<!-- grade_id: cell-0ecdfae26dbb43fd, points: 1, grade: true, solution: false -->
```python
# DO NOT MODIFY, THIS IS UNIT TESTING! [1 pt]
from nose.tools import assert_equal

assert_equal(toPalindromes(["cover", 123, "hi"]),
            {'cover': ('coverrevoc', 'revoccover'), 123: (123321, 321123), 'hi': ('hiih', 'ihhi')})
assert_equal(toPalindromes([235, "car"]),
             {235: (235532, 532235), 'car': ('carrac', 'raccar')})
assert_equal(toPalindromes(["and", "or", "not", 314]),
             {'and': ('anddna', 'dnaand'), 'or': ('orro', 'roor'), 'not': ('notton', 'tonnot'), 314: (314413, 413314)})
```

<!-- grade_id: cell-3336c707423a9f21, points: 1, grade: true, solution: false -->
```python
# HIDDEN TEST [1 pt]
```

3. Given the directory of information about students (dictionary) below, develop the following: (a) A function named `passing_gpa` that returns a dictionary where the keys are "ID" and the values are the "GPA" of students with GPA greater than or equal to the specified GPA threshold; (b) A function named `average_test_scores` that returns a dictionary where the keys are "ID" and the values are the average of test scores as a string (rounded off to two decimal places); (c) A function named `scholars_with_lastname` that returns a list of tuples which consists of "ID" and "Last Name" of students who has a last name and a scholar; (d) A function named `math_eng_sci` that returns a tuple of "ID" of students with all Math, English, and Science test scores greater than or equal to the specified score threshold.

<!-- grade_id: cell-f17ba2ccc79da39d, points: 0, grade: false, solution: true -->
```python
student_directory = {
    "Student 1" : {
        "ID" : "12198",
        "Last Name" : "Santa Cruz",
        "Gender" : 'Male',
        "GPA" : 3.4,
        "Scholar" : True,
        "Test Scores" : {"Math" : 90, "English" : 95, "Science" : 87}
    },
    "Student 2" : {
        "ID" : "12199",
        "Gender" : 'Male',
        "GPA" : 1.5,
        "Scholar" : True,
        "Test Scores" : {"Math" : 80, "English" : 75, "Science": 83}
    },
    "Student 3" : {
        "ID" : "12200",
        "Gender" : 'Female',
        "GPA" : 3.9,
        "Scholar" : False,
        "Test Scores" : {"Math" : 88, "English" : 85, "Science": 95}
    },
    "Student 4" : {
        "ID" : "12201",
        "Last Name" : "Delfino",
        "Gender" : 'Female',
        "GPA" : 1.4,
        "Scholar" : True,
        "Test Scores" : {"Math" : 82, "English" : 82, "Science" : 88}
    },
    "Student 5" : {
        "ID" : "12202",
        "Gender" : 'Female',
        "GPA" : 4.0,
        "Scholar" : True,
        "Test Scores" : {"Math" : 87, "English" : 91, "Science": 89}
    },
    "Student 6" : {
        "ID" : "12203",
        "Gender" : 'Male',
        "GPA" : 1.0,
        "Scholar" : False,
        "Test Scores" : {"Math" : 86, "English" : 92, "Science" : 94}
    },
    "Student 7" : {
        "ID" : "12204",
        "Last Name" : "Legarda",
        "Gender" : 'Female',
        "GPA" : 3.7,
        "Scholar" : True,
        "Test Scores" : {"Math" : 90, "English" : 95, "Science": 83}
    },
    "Student 8" : {
        "ID" : "12205",
        "Last Name" : "Ramos",
        "Gender" : 'Male',
        "GPA" : 2.8,
        "Scholar" : True,
        "Test Scores" : {"Math" : 80, "English" : 82, "Science" : 86}
    }
}

def passing_gpa(directory, gpa_thresh):
    if type(directory) is not dict:
        raise TypeError("First argument must be a dictionary")
    if type(gpa_thresh) is not float:
        raise TypeError("Second argument must be a float")
        
    result = {}
    for i in directory:
        if directory[i]["GPA"] >= gpa_thresh :
            result[directory[i]["ID"]]=directory[i]["GPA"]
    return result

def average_test_scores(directory):
    if type(directory) is not dict:
        raise TypeError("Argument must be a dictionary")
    result = {}
    for i in directory:
        ave = ((directory[i]["Test Scores"]["Math"]) + (directory[i]["Test Scores"]["English"]) + (directory[i]["Test Scores"]["Science"]))/3
        result[directory[i]["ID"]]="{:.2f}".format(ave)
    return result

def scholars_with_lastname(directory):
    if type(directory) is not dict:
        raise TypeError("Argument must be a dictionary")
    
    result = {}
    for i in directory:
        if directory[i]["Scholar"] == True and "Last Name" in directory[i]:
            result[i]=(directory[i]["ID"],directory[i]["Last Name"])
    resultList = [i for i in result.values()]
    return resultList

def math_eng_sci(directory, score_thresh):
    if type(directory) is not dict:
        raise TypeError("First argument must be a dictionary")
    if type(score_thresh) is not float:
        raise TypeError("Second argument must be a float")
    # YOUR CODE HERE
    resultDict = {}
    for i in directory:
        if directory[i]["Test Scores"]["Math"] >= score_thresh and directory[i]["Test Scores"]["English"] >= score_thresh and directory[i]["Test Scores"]["Science"] >= score_thresh:
            resultDict[i] = directory[i]["ID"]
    resultList = [i for i in resultDict.values()]
    return tuple(resultList)
```

<!-- grade_id: cell-861772d45b70f49d, points: 1, grade: true, solution: false -->
```python
# DO NOT MODIFY, THIS IS UNIT TESTING! [1 pt]
from nose.tools import assert_equal
from nose.tools import assert_is_instance

assert_equal(passing_gpa(student_directory, 1.5), 
             {'12198': 3.4, '12199': 1.5, '12200': 3.9, '12202': 4.0, '12204': 3.7, '12205': 2.8})
assert_equal(passing_gpa(student_directory, 3.0), 
             {'12198': 3.4, '12200': 3.9, '12202': 4.0, '12204': 3.7})
assert_equal(passing_gpa(student_directory, 3.5), 
             {'12200': 3.9, '12202': 4.0, '12204': 3.7})
assert_is_instance(passing_gpa(student_directory, 1.0), dict)
assert_is_instance(passing_gpa(student_directory, 1.0)['12198'], float)
```

<!-- grade_id: cell-1fa0ead94c035a39, points: 1, grade: true, solution: false -->
```python
# DO NOT MODIFY, THIS IS UNIT TESTING! [1 pt]
from nose.tools import assert_equal
from nose.tools import assert_is_instance

assert_equal(average_test_scores(student_directory),
             {'12198': '90.67', '12199': '79.33', '12200': '89.33', '12201': '84.00', 
              '12202': '89.00', '12203': '90.67', '12204': '89.33', '12205': '82.67'})
assert_equal(len(average_test_scores(student_directory)['12198']), 5)
assert_is_instance(average_test_scores(student_directory)['12200'], str)
```

<!-- grade_id: cell-5b133930567a30f6, points: 1, grade: true, solution: false -->
```python
# DO NOT MODIFY, THIS IS UNIT TESTING! [1 pt]
from nose.tools import assert_equal

assert_equal(scholars_with_lastname(student_directory),
             [('12198', 'Santa Cruz'), ('12201', 'Delfino'), ('12204', 'Legarda'), ('12205', 'Ramos')])
assert_is_instance(scholars_with_lastname(student_directory), list)
assert_is_instance(scholars_with_lastname(student_directory)[0], tuple)
```

<!-- grade_id: cell-a79777fa5fed871f, points: 1, grade: true, solution: false -->
```python
# DO NOT MODIFY, THIS IS UNIT TESTING! [1 pt]
from nose.tools import assert_equal
from nose.tools import assert_is_instance

assert_equal(math_eng_sci(student_directory, 80.0),
             ('12198', '12200', '12201', '12202', '12203', '12204', '12205'))
assert_equal(math_eng_sci(student_directory, 85.0),
             ('12198', '12200', '12202', '12203'))
assert_equal(math_eng_sci(student_directory, 87.0),
             ('12198', '12202'))
assert_is_instance(math_eng_sci(student_directory, 75.0), tuple)
assert_is_instance(math_eng_sci(student_directory, 75.0)[0], str)
```

4. Write a script that will allow Karel to translate a dictionary of tuples indicating the corners with colors to (a) $7 \times 7$ world and (b) $40 \times 40$ world. Each key of the dictionary pertains to a certain row and the elements of the tuple pertains to the columns of a particular row - a value of 1 indicates that the corner must be colored, and a value of 0 indicates that the corner must not be colored. You can choose any color as long as it is uniform. The dictionary for a $7 \times 7$ and $40 \times 40$ worlds are encoded in a separate code cell below.

    A single script should work for all world sizes. By convention, the first row and column starts at the most top-left corner of the world. If the supplied dictionary data is larger than the loaded world, it would simply ignore those data which cannot be placed on the world (e.g., data for columns 8 to 40 will be ignored in a $7 \times 7$ world).

    Two answer cells are provided below: (1) For the first cell, show a screenshot of the Karel application output for the  $7 \times 7$ world, (2) for the second cell, show a screenshot of the Karel application output for the  $40 \times 40$ world. Each cell is worth 1 point.

```python
# COPY THE FOLLOWING DICTIONARY DATA TO YOUR EXTERNAL SCRIPT FILE

# Dictionary data for a 7x7 world
pix_7x7 = {
    'Row 1': (0,0,0,0,0,0,0),
    'Row 2': (0,1,1,0,1,1,0),
    'Row 3': (0,1,1,0,1,1,0),
    'Row 4': (0,0,0,0,0,0,0),
    'Row 5': (0,1,0,0,0,1,0),
    'Row 6': (0,0,1,1,1,0,0),
    'Row 7': (0,0,0,0,0,0,0)
}

# Dictionary data for a 40x40 world
pix_40x40 = {
    'Row 1':  (0,0,0,0,0,0,0,0,0,0,1,0,0,0,1,0,0,0,0,0,0,0,0,0,0,0,1,0,0,0,1,0,0,0,0,0,0,0,0,0),
    'Row 2':  (0,0,0,0,0,0,0,0,0,1,0,0,0,1,0,0,0,0,0,0,0,0,0,0,0,0,0,1,0,0,0,1,0,0,0,0,0,0,0,0),
    'Row 3':  (0,0,0,0,0,0,0,0,1,1,0,0,1,1,0,0,0,0,0,0,0,0,0,0,0,0,0,1,1,0,0,1,1,0,0,0,0,0,0,0),
    'Row 4':  (0,0,0,0,0,0,0,1,1,0,0,1,1,0,0,0,0,0,0,0,0,0,0,0,0,0,0,0,1,1,0,0,1,1,0,0,0,0,0,0),
    'Row 5':  (0,0,0,0,0,0,0,1,0,0,1,1,0,0,0,0,0,0,0,0,0,0,0,0,0,0,0,0,0,1,1,0,0,1,0,0,0,0,0,0),
    'Row 6':  (0,0,0,0,0,0,1,1,0,0,1,1,0,0,0,0,0,0,0,0,0,0,0,0,0,0,0,0,0,1,1,0,0,1,1,0,0,0,0,0),
    'Row 7':  (0,0,0,0,0,0,1,1,0,1,1,0,0,0,0,0,0,0,0,0,0,0,0,0,0,0,0,0,0,0,1,1,0,1,1,0,0,0,0,0),
    'Row 8':  (0,0,0,0,0,1,1,1,0,1,1,0,0,0,0,0,0,0,0,0,0,0,0,0,0,0,0,0,0,0,1,1,0,1,1,1,0,0,0,0),
    'Row 9':  (0,0,0,0,0,1,1,0,0,1,1,0,0,0,0,0,0,0,0,0,0,0,0,0,0,0,0,0,0,0,1,1,0,0,1,1,0,0,0,0),
    'Row 10': (0,0,0,0,0,1,1,0,1,1,1,0,0,0,0,0,0,0,0,0,0,0,0,0,0,0,0,0,0,0,1,1,1,0,1,1,0,0,0,0),
    'Row 11': (0,0,0,0,0,1,1,0,1,1,1,1,1,1,1,1,0,0,0,1,0,1,0,0,0,1,1,1,1,1,1,1,1,0,1,1,0,0,0,0),
    'Row 12': (0,0,0,0,0,1,1,0,1,1,1,1,1,1,1,1,1,0,1,0,0,0,1,0,1,1,1,1,1,1,1,1,1,0,1,1,0,0,0,0),
    'Row 13': (0,0,0,0,0,1,1,0,0,0,0,0,0,1,1,0,0,1,1,1,1,1,1,1,0,0,1,1,0,0,0,0,0,0,1,1,0,0,0,0),
    'Row 14': (0,0,0,0,0,1,1,0,0,0,1,1,1,1,1,1,0,1,1,1,1,1,1,1,0,1,1,1,1,1,1,0,0,0,1,1,0,0,0,0),
    'Row 15': (0,0,0,0,0,0,1,1,1,1,1,1,1,0,0,1,1,1,1,1,1,1,1,1,1,1,0,0,1,1,1,1,1,1,1,0,0,0,0,0),
    'Row 16': (0,0,0,0,0,0,0,1,1,1,1,0,0,0,0,0,0,1,1,1,1,1,1,1,0,0,0,0,0,0,1,1,1,1,0,0,0,0,0,0),
    'Row 17': (0,0,0,0,0,0,0,0,0,0,0,0,0,0,1,1,1,1,1,1,1,1,1,1,1,1,1,0,0,0,0,0,0,0,0,0,0,0,0,0),
    'Row 18': (0,0,0,0,0,0,0,0,0,0,0,0,0,1,1,0,0,1,1,1,1,1,1,1,0,0,1,1,0,0,0,0,0,0,0,0,0,0,0,0),
    'Row 19': (0,0,0,0,0,0,0,0,0,0,0,0,1,1,0,0,1,1,1,1,1,1,1,1,1,0,0,1,1,0,0,0,0,0,0,0,0,0,0,0),
    'Row 20': (0,0,0,0,0,0,0,0,0,0,0,1,1,0,0,1,1,0,0,1,1,1,0,0,1,1,0,0,1,1,0,0,0,0,0,0,0,0,0,0),
    'Row 21': (0,0,0,0,0,0,0,0,0,1,1,1,0,1,0,0,0,1,1,1,1,1,1,1,0,0,0,1,0,1,1,1,0,0,0,0,0,0,0,0),
    'Row 22': (0,0,0,0,0,1,0,0,1,1,1,0,0,1,1,0,0,1,1,1,1,1,1,1,0,0,1,1,0,0,1,1,1,0,0,1,0,0,0,0),
    'Row 23': (0,0,0,0,0,1,1,1,1,0,0,0,0,1,1,0,0,1,1,1,1,1,1,1,0,0,1,1,0,0,0,0,1,1,1,1,0,0,0,0),
    'Row 24': (0,0,0,0,0,1,1,1,0,0,0,0,0,1,1,0,0,1,1,1,1,1,1,1,0,0,1,1,0,0,0,0,0,1,1,1,0,0,0,0),
    'Row 25': (0,0,0,0,0,0,1,1,0,0,0,0,0,1,1,0,0,0,1,1,1,1,1,0,0,0,1,1,0,0,0,0,0,1,1,0,0,0,0,0),
    'Row 26': (0,0,0,0,0,0,1,1,0,0,0,0,0,1,1,0,0,0,1,1,1,1,1,0,0,0,1,1,0,0,0,0,0,1,1,0,0,0,0,0),
    'Row 27': (0,0,0,0,0,0,1,1,0,0,0,0,0,1,1,0,0,0,1,1,1,1,1,0,0,0,1,1,0,0,0,0,0,1,1,0,0,0,0,0),
    'Row 28': (0,0,0,0,0,0,1,1,0,0,0,0,0,1,1,0,0,0,1,1,1,1,1,0,0,0,1,1,0,0,0,0,0,1,1,0,0,0,0,0),
    'Row 29': (0,0,0,0,0,0,1,1,0,0,0,0,0,1,1,0,0,0,0,1,1,1,0,0,0,0,1,1,0,0,0,0,0,1,1,0,0,0,0,0),
    'Row 30': (0,0,0,0,0,0,1,1,0,0,0,0,0,1,1,0,0,0,0,1,0,1,0,0,0,0,1,1,0,0,0,0,0,1,1,0,0,0,0,0),
    'Row 31': (0,0,0,0,0,0,0,1,1,0,0,0,0,1,1,0,0,0,0,0,0,0,0,0,0,0,1,1,0,0,0,0,1,1,0,0,0,0,0,0),
    'Row 32': (0,0,0,0,0,0,0,1,1,0,0,0,0,0,1,1,0,0,0,0,0,0,0,0,0,1,1,0,0,0,0,0,1,1,0,0,0,0,0,0),
    'Row 33': (0,0,0,0,0,0,0,1,1,0,0,0,0,0,1,1,0,0,0,0,0,0,0,0,0,1,1,0,0,0,0,0,1,1,0,0,0,0,0,0),
    'Row 34': (0,0,0,0,0,0,0,1,1,0,0,0,0,0,1,1,0,0,0,0,0,0,0,0,0,1,1,0,0,0,0,0,1,1,0,0,0,0,0,0),
    'Row 35': (0,0,0,0,0,0,0,0,1,1,0,0,0,0,0,1,0,0,0,0,0,0,0,0,0,1,0,0,0,0,0,1,1,0,0,0,0,0,0,0),
    'Row 36': (0,0,0,0,0,0,0,0,0,1,0,0,0,0,0,1,1,0,0,0,0,0,0,0,1,1,0,0,0,0,0,1,0,0,0,0,0,0,0,0),
    'Row 37': (0,0,0,0,0,0,0,0,0,0,1,0,0,0,0,0,1,1,0,0,0,0,0,1,1,0,0,0,0,0,1,0,0,0,0,0,0,0,0,0),
    'Row 38': (0,0,0,0,0,0,0,0,0,0,0,1,0,0,0,0,0,1,0,0,0,0,0,1,0,0,0,0,0,1,0,0,0,0,0,0,0,0,0,0),
    'Row 39': (0,0,0,0,0,0,0,0,0,0,0,0,0,0,0,0,0,1,0,0,0,0,0,1,0,0,0,0,0,0,0,0,0,0,0,0,0,0,0,0),
    'Row 40': (0,0,0,0,0,0,0,0,0,0,0,0,0,0,0,0,0,0,1,0,0,0,1,0,0,0,0,0,0,0,0,0,0,0,0,0,0,0,0,0)
}
```

<!-- grade_id: cell-349cf941f53f1f52, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: cell-b88b738a393eb33d, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

***

<!-- grade_id: cell-4420f5b4060accd6, points: 0, grade: true, solution: false -->
```python

```