# Programming Logic and Design Laboratory (LBYCPA1)

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or `YOUR ANSWER HERE`, as well as your name and other information below inside the double quotes. Remove the line `raise NotImplementedError()` once you have answered the problem.

REMINDERS: Section 4.15 of the Student Handbook (SH) 2018-2021 provides that:

As a Catholic University that prepares its students to become professionals and be the resource for Church and nation, De La Salle University puts a premium on academic honesty. A student's academic requirements, such as assignments, term papers, computer programs/projects, and thesis papers, should be his/her own work. He/She must distinguish between his/her own ideas and those of other authors.

The student must cite references, direct quotes, and other sources (including data obtained from tables, illustrations, figures, pictures, images, and video) following the prescribed format of the discipline. Should a student intend to submit a similar previous work for credit, in different courses, a written permission must first be obtained from the faculty members concerned. In addition, for works where he/she is only one of the authors, the following must be observed.

4.15.1 Provide a disclaimer stating that the work is not entirely his/hers;

4.15.2 The work cannot be published without written consent of the group of authors.

4.15.3 If the work is separable (i.e. parts are clearly attributed to the individual authors), the student may use and must cite his/her own part only; otherwise, the student must seek the written consent of the group of authors.

4.15.4 Computer codes obtained from open source reference may be copied freely, but the source should be cited, provided further that the faculty member allows the use of open source reference in the course.

## Rubric for the Final Laboratory Report

|Criteria|EXEMPLARY (90-100)|SATISFACTORY (80-89)|DEVELOPING (70-79)|BEGINNING (Below 70)|WEIGHT|
|:-|:-|:-|:-|:-|:-:|
|**Experimental Plan (Flowchart/ Algorithm)**|Experimental plan has supporting details and diagram/algorithm that is stated and well explained|Experimental plan has supporting details and diagram/algorithm that is stated but not explained|Experimental plan is vague or brief. It has supporting details and doesn’t have diagram/algorithm|No experimental plan presented|30%|
|**Conducting Experiment/ Activity**|Objective and Materials used (modern engineering tools, software, and instruments/equipment) are identified. Steps are easy to follow for conducting and experiment/activity|Objective is not stated. Materials used (modern engineering tools, software, and instruments/equipment) are identified. Steps are easy to follow for conducting and experiment/activity|Does not provide enough information to conduct an experiment/activity|No Objective, Materials used (modern engineering tools, software, and instruments/equipment), and steps for conducting experiment/activity provided|20%|
|**Codes/Data/ Program**|Data is well utilized in the program. Program codes are easy to read. Program output has no error. Questions are answered completely and correctly|Data is somewhat utilized in the program. Program code are easy to read. Program output has an output but logically incorrect. Some questions are answered completely and correctly|Data is not utilized in the program. It has a missing significant code/syntax in the program|No program presented|30%|
|**Grammar, logical presentation, and format**|The report was grammatically correct, logically presented and used the required format|The report had minimal grammatical errors and somewhat presented logically. The required format was used|The report had a lot of grammatical errors and not logically presented; the required format was barely used|The report had a lot of grammatical errors, was not logically presented and the required format was not used|20%|

```python
# Fill-up the following information (put inside double quotes)

NAME = "DELA CRUZ, JOHN CARLO THEO S."                   # e.g. LIGUTAN, DINO DOMINIC F.
STUDENT_NUMBER = "12195758"         # e.g. 20192118
SECTION = "EQ1"                # e.g. EQ1
DATE_TIME_COMPLETED = "06/22/22, 8:23pm"    # e.g. 03/21/2022,11:00
```

<hr style="border:1px solid black"> </hr>

# Problem Set 10
<hr style="border:1px solid gray"> </hr>

1. Implement a function that will accept one or more numbers (a vector) and returns the norm of the vector (format as a string, with two decimal places). The norm of vector ($\left|\mathbf{v}\right|$) is defined as
<center>$$\left|\mathbf{v}\right|=\sqrt{\sum\limits_{i=1}^n x_i}=\sqrt{x_1^2+x_2^2+...+x_n^2}$$</center>
    Provide exception handling for any non-numerical input by returning the string "An input is not a valid number"

<!-- grade_id: cell-32216f35a45b784e, points: 0, grade: false, solution: true -->
```python
from math import sqrt
def norm(num1, *nums): 
    
    numbers = 0
    if type(num1) is str:
        return "An input is not a valid number"
    for i in nums:
        if type(i) is str:
            return "An input is not a valid number"    
    if nums:
        for i in nums:
            square = i**2
            numbers += square
    vector = sqrt((num1**2) + numbers)
    
    return "{:.2f}".format(vector)
```

<!-- grade_id: cell-a863d6410ce95310, points: 1, grade: true, solution: false -->
```python
# DO NOT MODIFY, THIS IS UNIT TESTING! [1 pt]
from nose.tools import assert_equal

assert_equal(norm(0), '0.00')
assert_equal(norm(4, 6, 7), '10.05')
assert_equal(norm(0.5, 0.8, 0.4, 0.5, 0.7, 0.6), '1.47')
assert_equal(norm(2, '6', 8, 7, 9), 'An input is not a valid number')
num_list = (5.3, 8.1, 2.4, 5.6, 7.7)
assert_equal(norm(*num_list), '13.79')
```

<!-- grade_id: cell-b25b3cbf30198389, points: 1, grade: true, solution: false -->
```python
# HIDDEN TEST [1 pt]
```

2. Develop a function named `cycle_permute` that will perform a cyclic permutation on an arbitrary number of arguments. Input arguments could either be a numeric or string type. It should have keyword argument named `steps` (an integer), the number of elements to shift the cyclic permutation. A positive integer would shift the elements to the right and a negative integer would shift the elements to the left. The default value `steps=0` will return the same sequence. The function must return a tuple of permuted arguments. See unit testing cell below for examples.

<!-- grade_id: cell-2a67d9bb9ce95749, points: 0, grade: false, solution: true -->
```python
def cycle_permute(*nums, steps=0):
    temp = list(nums)
    if steps < 0:
        for i in range(0, steps, -1):
            removeElement = temp.pop(0)
            temp.append(removeElement)
    else:
        for i in range(steps):
            removeElement = temp.pop(-1)
            temp.insert(0,removeElement)

    return tuple(temp)
```

<!-- grade_id: cell-0ecdfae26dbb43fd, points: 1, grade: true, solution: false -->
```python
# DO NOT MODIFY, THIS IS UNIT TESTING! [1 pt]
from nose.tools import assert_equal

assert_equal(cycle_permute(1, 2, 3.0, 4, '5', '6.0', 7, steps=-1), (2, 3.0, 4, '5', '6.0', 7, 1))
assert_equal(cycle_permute('go', 'glow', 'grow'), ('go', 'glow', 'grow'))
assert_equal(cycle_permute(0.4, 'i', '<3', 8, 'u', steps=8), ('<3', 8, 'u', 0.4, 'i'))
arg_list = ['I', '<3', 'programming', 3000]
assert_equal(cycle_permute(*arg_list, steps=5), (3000, 'I', '<3', 'programming'))
arg_dict = {'we': 2, 'are': 3, 'beginners': 9, 1:0}
assert_equal(cycle_permute(*arg_dict, steps=3), ('are', 'beginners', 1, 'we'))
```

<!-- grade_id: cell-3336c707423a9f21, points: 1, grade: true, solution: false -->
```python
# HIDDEN TEST [1 pt]
```

3. Write a function named `updateInventory` that will accept a filename followed by an arbitrary number of keyword arguments. The file is in comma-separated values (csv) format that contains a list of inventory items and quantity. The function should update the contents of the specified csv file based from the supplied keyword arguments. The first argument (`filename`) is a string containing the path of the file and the rest of the keyword arguments specifies the inventory item name and the corresponding quantity as the value.

    Each line in the csv file corresponds to an inventory item and its quantity. The first column contains the inventory item name and the second column contains the quantity. Supplying a keyword argument to the function that matches an existing inventory item replaces its quantity by the assigned value. If a supplied keyword argument is new, it will be added as a new inventory item in a new line to the csv file. If no keyword arguments were supplied, nothing will be updated to the file. Inventory item names must not contain spaces. See unit testing cell below for examples.
    
    The function must return a dictionary of updated inventory item names (`str`) and quantity (`int`). You can use the `map` function to generate the output dictionary. Provide exception handling if the file cannot be accessed by returning the string "The specified file cannot be accessed"

<!-- grade_id: cell-ea42ec0c524af7ef, points: 0, grade: false, solution: true -->
```python
import csv
def updateInventory(filename, **parts):
    dk = {}  
    with open(filename, "r") as inventory_file: 
        f = inventory_file.read()
        f = f.strip()
        f = f.replace('\n',',')
        f = f.split(',')
    
        for i in range(0, len(f), 2):
            dk[f[i]] = int(f[i+1])
        dk.update(parts)
        
    with open(filename, "w")  as inventory_file:
        for i in dk:
            inventory_file.write("{},{}\n".format(i, dk[i]))
        
    return dk        
```

```python
updateInventory("inventory.csv", SSD_1TB=7, HDD_2TB=15)
```

Output:
```
{'CPU_AMD_Ryzen': 6,
 'NVIDIA_RTX3060': 2,
 'RAM_4GB': 10,
 'HDD_2TB': 15,
 'CPU_Intel_i9': 3,
 'SSD_1TB': 7}
```

<!-- grade_id: cell-d4a5f354ec603bc4, points: 1, grade: true, solution: false -->
```python
# DO NOT MODIFY, THIS IS UNIT TESTING! [1 pt]
from nose.tools import assert_equal

with open("inventory.csv", "w") as inventory_file: # create a test file
    inventory_file.write("CPU_Intel_i9,5\nHDD_2TB,15")

assert_equal(updateInventory("inventory.csv", RAM_4GB=7), 
             {'CPU_Intel_i9': 5, 'HDD_2TB': 15, 'RAM_4GB': 7})
assert_equal(updateInventory("inventory.csv", SSD_1TB=7, HDD_2TB=15),
             {'CPU_Intel_i9': 5, 'HDD_2TB': 15, 'RAM_4GB': 7, 'SSD_1TB': 7})
assert_equal(updateInventory("inventory.csv", RAM_2GB=12, CPU_Intel_i9=3, SSD_1TB=5),
             {'CPU_Intel_i9': 3, 'HDD_2TB': 15, 'RAM_4GB': 7, 'SSD_1TB': 5, 'RAM_2GB': 12})
```

<!-- grade_id: cell-a7cb5b583a8cc0e9, points: 1, grade: true, solution: false -->
```python
# DO NOT MODIFY, THIS IS UNIT TESTING! [1 pt]
from nose.tools import assert_equal

with open("inventory.csv", "w") as inventory_file: # create a test file
    inventory_file.write("CPU_AMD_Ryzen,6\nNVIDIA_RTX3060,3\nRAM_4GB,10")
updateInventory("inventory.csv", HDD_2TB=6, CPU_Intel_i9=3, NVIDIA_RTX3060=2)

with open("inventory.csv") as inventory_file: # read the test file
    items = inventory_file.read().split("\n")
assert_equal(int(next(item for item in items if 'CPU_AMD_Ryzen' in item).split(',')[1]), 6)
assert_equal(int(next(item for item in items if 'NVIDIA_RTX3060' in item).split(',')[1]), 2)
assert_equal(int(next(item for item in items if 'RAM_4GB' in item).split(',')[1]), 10)
assert_equal(int(next(item for item in items if 'HDD_2TB' in item).split(',')[1]), 6)
assert_equal(int(next(item for item in items if 'CPU_Intel_i9' in item).split(',')[1]), 3)
```

<!-- grade_id: cell-d8aebd5ab92f3e12, points: 1, grade: true, solution: false -->
```python
# HIDDEN TEST [1 pt]
```

4. Given an arbitrary sentence, implement a function named `wordsWithRepeatingLetters` that will return a set of words such that each word contains at least a letter that repeats at least twice. Consider the uppercase and lowercase of a letter to be the same symbols. Separate by whitespace only and remove all periods (`.`) prior to processing. You can use the `filter` and `any` functions to facilitate the output of such words as a set. Use of conditional (e.g., `if`, etc.) and looping (e.g., `for`, etc.) statements are not allowed for this problem.

<!-- grade_id: cell-f17ba2ccc79da39d, points: 0, grade: false, solution: true -->
```python
def wordsWithRepeatingLetters(sentence):
    words = list(map(lambda x: x.strip('.'), sentence.split()))
    reference = set(map(lambda x: x, sentence.lower())) #Alphabet
    filtered = set(filter(lambda word:(any(map(lambda y : word.lower().count(y) >= 2, reference ))),words))
    return filtered
```

```python
# def wordsWithRepeatingLetters(sentence):
#     return set(filter(lambda word:(any(map(lambda y : word.lower().count(y) >= 2, set(map(lambda x: x, sentence.lower())) ))), list(map(lambda x: x.strip('.'), sentence.split()))))
```

<!-- grade_id: cell-861772d45b70f49d, points: 1, grade: true, solution: false -->
```python
# DO NOT MODIFY, THIS IS UNIT TESTING! [1 pt]
from nose.tools import ok_, assert_is_instance

sentence = "He heard the loud impact before he ever saw the result. "\
"It had been so loud that it had actually made him jump back in his seat. "\
"As soon as he recovered from the surprise, he saw the crack in the windshield. "\
"It seemed to be an analogy of the current condition of his life."

words_with_repeating_letters_1 = wordsWithRepeatingLetters(sentence)
ok_("condition" in words_with_repeating_letters_1) # since there are two i's
ok_('actually' in words_with_repeating_letters_1) # since there are two a's or l's
ok_('been' in words_with_repeating_letters_1) # since there are two e's
ok_('that' in words_with_repeating_letters_1) # since there are two t's
ok_('current' in words_with_repeating_letters_1) # since there are two r's
assert_is_instance(words_with_repeating_letters_1, set)
```

<!-- grade_id: cell-2ed873f693a2c988, points: 1, grade: true, solution: false -->
```python
# DO NOT MODIFY, THIS IS UNIT TESTING! [1 pt]
from nose.tools import ok_, assert_is_instance

sentence = "They rushed out the door, grabbing anything and everything they could think of they might need. "\
"There was no time to double-check to make sure they weren't leaving something important behind. "\
"Everything was thrown into the car and they sped off. "\
"Thirty minutes later they were safe and that was when it dawned on them that they had forgotten the most important thing of all."

words_with_repeating_letters_2 = wordsWithRepeatingLetters(sentence)
ok_("Thirty" in words_with_repeating_letters_2) # since there are two t's
ok_("all" in words_with_repeating_letters_2) # since there are two l's
ok_("dawned" in words_with_repeating_letters_2) # since there are two d's
ok_("grabbing" in words_with_repeating_letters_2) # since there are two g's or b's
ok_("need" in words_with_repeating_letters_2) # since there are two e's
assert_is_instance(words_with_repeating_letters_2, set)
```

<!-- grade_id: cell-a79777fa5fed871f, points: 1, grade: true, solution: false -->
```python
# HIDDEN TEST [1 pt]
```

***

<!-- grade_id: cell-4420f5b4060accd6, points: 0, grade: true, solution: false -->
```python

```