# Programming Logic and Design Laboratory (LBYCPA1)

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or `YOUR ANSWER HERE`, as well as your name and other information below inside the double quotes. Remove the line `raise NotImplementedError()` once you have answered the problem.

REMINDERS: Section 4.15 of the Student Handbook (SH) 2018-2021 provides that:

As a Catholic University that prepares its students to become professionals and be the resource for Church and nation, De La Salle University puts a premium on academic honesty. A student's academic requirements, such as assignments, term papers, computer programs/projects, and thesis papers, should be his/her own work. He/She must distinguish between his/her own ideas and those of other authors.

The student must cite references, direct quotes, and other sources (including data obtained from tables, illustrations, figures, pictures, images, and video) following the prescribed format of the discipline. Should a student intend to submit a similar previous work for credit, in different courses, a written permission must first be obtained from the faculty members concerned. In addition, for works where he/she is only one of the authors, the following must be observed.

4.15.1 Provide a disclaimer stating that the work is not entirely his/hers;

4.15.2 The work cannot be published without written consent of the group of authors.

4.15.3 If the work is separable (i.e. parts are clearly attributed to the individual authors), the student may use and must cite his/her own part only; otherwise, the student must seek the written consent of the group of authors.

4.15.4 Computer codes obtained from open source reference may be copied freely, but the source should be cited, provided further that the faculty member allows the use of open source reference in the course.

## Rubric for the Final Laboratory Report

|Criteria|EXEMPLARY (90-100)|SATISFACTORY (80-89)|DEVELOPING (70-79)|BEGINNING (Below 70)|WEIGHT|
|:-|:-|:-|:-|:-|:-:|
|**Experimental Plan (Flowchart/ Algorithm)**|Experimental plan has supporting details and diagram/algorithm that is stated and well explained|Experimental plan has supporting details and diagram/algorithm that is stated but not explained|Experimental plan is vague or brief. It has supporting details and doesn’t have diagram/algorithm|No experimental plan presented|30%|
|**Conducting Experiment/ Activity**|Objective and Materials used (modern engineering tools, software, and instruments/equipment) are identified. Steps are easy to follow for conducting and experiment/activity|Objective is not stated. Materials used (modern engineering tools, software, and instruments/equipment) are identified. Steps are easy to follow for conducting and experiment/activity|Does not provide enough information to conduct an experiment/activity|No Objective, Materials used (modern engineering tools, software, and instruments/equipment), and steps for conducting experiment/activity provided|20%|
|**Codes/Data/ Program**|Data is well utilized in the program. Program codes are easy to read. Program output has no error. Questions are answered completely and correctly|Data is somewhat utilized in the program. Program code are easy to read. Program output has an output but logically incorrect. Some questions are answered completely and correctly|Data is not utilized in the program. It has a missing significant code/syntax in the program|No program presented|30%|
|**Grammar, logical presentation, and format**|The report was grammatically correct, logically presented and used the required format|The report had minimal grammatical errors and somewhat presented logically. The required format was used|The report had a lot of grammatical errors and not logically presented; the required format was barely used|The report had a lot of grammatical errors, was not logically presented and the required format was not used|20%|

```python
# Fill-up the following information (put inside double quotes)

NAME = "DELA CRUZ, JOHN CARLO THEO S."                 
STUDENT_NUMBER = "12195758"        
SECTION = "EQ1"               
DATE_TIME_COMPLETED = "04/11/2022, 10:53PM"  
```

<hr style="border:1px solid black"> </hr>

# Problem Set 3
<hr style="border:1px solid gray"> </hr>

1. Given the grading scheme displayed in the table below (with 70 passing mark), write a function that converts a raw grade $x$,  $0.0 \le x \le 100.0$ to its corresponding DLSU grade equivalent, $\{0.0, 1.0, ..., 3.5, 4.0\}$.

| Raw Score | Grade Equivalent |
| --------- | ---------------- |
| 100% down to 96% | 4.0 |
| Less than 96% down to 92% | 3.5 |
| Less than 92% down to 88% | 3.0 |
| Less than 88% down to 83% | 2.5 |
| Less than 83% down to 77% | 2.0 |
| Less than 77% down to 74% | 1.5 |
| Less than 74% down to 70% | 1.0 |
| Less than 70% | 0.0 |

<!-- grade_id: cell-32216f35a45b784e, points: 0, grade: false, solution: true -->
```python
def raw2grade(score):
    # This function converts a raw score to DLSU grade equivalent (70 passing score)
    
    if not (type(score) is int or type(score) is float): #It verifies a correct input 
        raise TypeError("Input not a valid number")
            
    if score > 100 or score < 0:
        raise ValueError("Raw score must not be greater than 100 nor less than 0")

    if score >= 96 and score <= 100:
        gpa = 4.0
    elif score >= 92 and score < 96:
        gpa = 3.5
    elif score >= 88 and score < 92:
        gpa = 3.0
    elif score >= 83 and score < 88:
        gpa = 2.5
    elif score >= 77 and score < 83:
        gpa = 2.0
    elif score >= 74 and score < 77:
        gpa = 1.5
    elif score >= 70 and score < 74:
        gpa = 1.0
    else:
        gpa = 0.0
       
    return gpa # Make sure to assign the result to the gpa variable

score = float(input("DLSU GPA Raw Grade Grade Calculator: "))
gpa = raw2grade(score)
print("GPA:", gpa)
```

Output:
```
DLSU GPA Raw Grade Grade Calculator: 90
GPA: 3.0
```

<!-- grade_id: cell-a863d6410ce95310, points: 2, grade: true, solution: false -->
```python
# DO NOT MODIFY, THIS IS UNIT TESTING!

assert raw2grade(97) == 4.0
assert raw2grade(93) == 3.5
assert raw2grade(89) == 3.0
assert raw2grade(85) == 2.5
assert raw2grade(80) == 2.0
assert raw2grade(75) == 1.5
assert raw2grade(70) == 1.0
assert raw2grade(65) == 0.0

# HIDDEN TEST
```

```python
# SAMPLE TEST, DO NOT MODIFY! JUST EXECUTE ONLY FOR OUTPUT
# You can add a cell below if you need to test a different input

print(raw2grade(96.5))
print(raw2grade(83.1))
```

Output:
```
4.0
2.5
```

2. In Spherical Trigonometry, the cosine of any side of a spherical triangle is equal to the product of  the cosines of the remaining two sides added to the product of the sines of those two sides multiplied by the cosine    of their included angle:
$$
\cos a = \cos b \cos c +  \sin b \sin c \cos A
$$
For the angles, the formula becomes:
$$
\cos A = -\cos B \cos C +  \sin B \sin C \cos a
$$
Write a function that computes the side or the angle of a spherical triangle depending on a mode variable, e.g., mode can be "side" (default) or "angle".


<!-- grade_id: cell-2a67d9bb9ce95749, points: 0, grade: false, solution: true -->
```python
def trig_compute(bB, cC, Aa, mode="side"):
    # This function computes the side or the angle depending on a mode variable, e.g. mode can be "side" or "angle". 
    
    from math import sin, cos, acos # We will use sin, cos, and inverse cos functions
    
    # YOUR CODE HERE
    if mode == "side" or mode == "SIDE":
        answer = acos(cos(bB)*cos(cC) + sin(bB)*sin(cC)*cos(Aa))
    elif mode == "angle" or mode == "ANGLE":
        answer = acos(-cos(bB)*cos(cC)+ sin(bB)*sin(cC)*cos(Aa))
    else:
        print("! Error !")
    return answer # Make sure to assign the result to the answer variable
mode = str(input("Hello \nSide or Angle?: "))
Aa = float(input("Input Value A: "))
bB = float(input("Input Value B: "))
cC = float(input("Input Value C: "))

answer = trig_compute(bB, cC, Aa, mode="side")
print(answer)
```

Output:
```
Hello 
Side or Angle?: side
Input Value A: 1
Input Value B: 2
Input Value C: 3
1.0686433167004479
```

<!-- grade_id: cell-0ecdfae26dbb43fd, points: 2, grade: true, solution: false -->
```python
# DO NOT MODIFY, THIS IS UNIT TESTING!
from math import pi, acos, sqrt

assert abs(trig_compute(pi/3, pi/6, pi/2) - acos(sqrt(3)/4)) < 1e-8
assert abs(trig_compute(pi/3, pi/6, pi/2, mode="angle") - acos(-sqrt(3)/4)) < 1e-8

# HIDDEN TEST
```

```python
# SAMPLE TEST, DO NOT MODIFY! JUST EXECUTE ONLY FOR OUTPUT
# You can add a cell below if you need to test a different input

print(trig_compute(pi/3, pi/6, pi/2))
print(trig_compute(pi/3, pi/6, pi/2, mode="angle"))
```

Output:
```
1.122963929865964
2.0186287237238294
```

3. Let $p$ and $q$ be propositions (a statement that can either be True or False). The conditional statement $p → q$ is the proposition "if $p$, then $q$." The conditional statement $p → q$ is false when $p$ is true and $q$ is false, and true otherwise. Write a function that implements the conditional logical operation.

<!-- grade_id: cell-5bf885b89699825d, points: 0, grade: false, solution: true -->
```python
def conditional(p, q):
    # This function implements the conditional logical operation
    
    if not (type(p) is bool and type(q) is bool):
        raise TypeError("Inputs are not valid Booleans!")
        
    if p is True and q is False:
        evaluation = False
    else:
        evaluation = True
    
    return evaluation # Make sure to assign the result to the evaluation variable

print("1 = True, 0 = False\n\n")

p = bool(int(input("Input P: ")))
q = bool(int(input("Input Q: ")))
evaluation = conditional(p,q)
print("The statement is : ", evaluation)
```

Output:
```
1 = True, 0 = False


Input P: 1
Input Q: 0
The statement is :  False
```

<!-- grade_id: cell-67877fb08827f061, points: 2, grade: true, solution: false -->
```python
# DO NOT MODIFY, THIS IS UNIT TESTING!

assert conditional(True, False) == False
assert conditional(True, True) == True

# HIDDEN TEST
```

```python
# SAMPLE TEST, DO NOT MODIFY! JUST EXECUTE ONLY FOR OUTPUT

print(conditional(False, False))
print(conditional(False, True))
print(conditional(True, False))
print(conditional(True, True))
```

Output:
```
True
True
False
True
```

4. The 8-digit I.D. numbers in a certain university can be verified by multiplying the 1st digit by 8, the 2nd digit by 7, and so on, until the last digit is multiplied by 1. Then, the sum of the series of products should be divisible by 11 to be a valid I.D. number; otherwise, the I.D. number is invalid.
```
Ex. 1.  I.D. No.: 11191031 :     
1*8 + 1*7 + 1*6 + 9*5 + 1*4 + 0*3 + 3*2 + 1*1 equals 77
Since the number 77 is divisible by 11, the I.D. is valid!
```
```
Ex. 2.  I.D. No.: 11191013 :
1*8 + 1*7 + 1*6 + 9*5 + 1*4 + 0*3 + 1*2 + 3*1 equals 75 
Since the number 75 is NOT divisible by 11, the I.D. is invalid!
```
Suppose that the resulting sum when divided by 11 is less than 16 for a Student I.D. and greater or equal to 16 for Faculty I.D. for that particular university. Implement a Python program that distinguishes the Student from Faculty I.D. Numbers among the valid I.D. numbers. The program should output the strings "INVALID", "STUDENT", and "FACULTY", respectively.

<!-- grade_id: cell-ddf77fa3ab60f546, points: 0, grade: false, solution: true -->
```python
def id_verify(num):
    # This function returns the strings "INVALID", "STUDENT", and "FACULTY" for each corresponding ID category
    
    if type(num) is not int:
        raise TypeError("Input not a valid integer")
        
    dig1 = (num % 100000000 // 10000000)
    dig2 = (num % 10000000 // 1000000) 
    dig3 = (num % 1000000 // 100000)
    dig4 = (num % 100000 // 10000) 
    dig5 = (num % 10000 // 1000)
    dig6 = (num % 1000 // 100) 
    dig7 = (num % 100 // 10) 
    dig8 = (num % 10) 
    total = (dig1 * 8) + (dig2 * 7) + (dig3 * 6) + (dig4 * 5) + (dig5 * 4) + (dig6 * 3) + (dig7 * 2) + (dig8 * 1)
    
    if(total % 11) == 0:  
        if(total / 11) < 16:
            identity = "STUDENT"
        else: 
            identity = "FACULTY"       
    else:
        identity = "INVALID"
    
    return identity # Make sure to assign the result to the identity variable

num = int(input("Enter your ID Number: "))
identity = id_verify(num) #declaring the formulated function
print(identity)
```

Output:
```
Enter your ID Number: 12195758
STUDENT
```

<!-- grade_id: cell-9fb1a74bdba298aa, points: 2, grade: true, solution: false -->
```python
# DO NOT MODIFY, THIS IS UNIT TESTING!

assert id_verify(12191031) == "INVALID"
assert id_verify(11843691) == "STUDENT"
assert id_verify(12191013) == "INVALID"
assert id_verify(11817054) == "STUDENT"
assert id_verify(11720721) == "INVALID"
assert id_verify(11831790) == "STUDENT"
assert id_verify(11831709) == "INVALID"
assert id_verify(11810661) == "STUDENT"
assert id_verify(11810616) == "INVALID"
assert id_verify(98078305) == "FACULTY"

# HIDDEN TEST
```

```python
# SAMPLE TEST, DO NOT MODIFY, JUST EXECUTE ONLY FOR OUTPUT
# You can add a cell below if you need to test a different input

id_num = int(input("Input your ID number: "))
print(id_verify(id_num))
```

Output:
```
Input your ID number: 12195758
STUDENT
```

5. Create a Python code that will check if two given strings are rotations of each other. Hint: You can use the `in` operator in a conditional statement.

<!-- grade_id: cell-f17ba2ccc79da39d, points: 0, grade: false, solution: true -->
```python
def string_rotations(str1, str2):
    # This function checks if the given two strings are rotations of each other
    # Hint: You can use the 'in' operator and string concatenation
    
    if not (type(str1) is str and type(str2) is str):
        raise TypeError("Inputs are not a valid strings!")
        
    # YOUR CODE HERE
    if str2 in str1 + str1:
        evaluation = True
    else:
        evaluation = False
    
    return evaluation # Make sure to assign the result to the evaluation variable

str1 = (input("Input a series of letters and/or numbers: "))
str2 = (input("Input 2nd Letter/Number: "))

evaluation = string_rotations(str1, str2)

print("Rotation: ", evaluation)
```

Output:
```
Input a series of letters and/or numbers: abacd
Input 2nd Letter/Number: cdaba
Rotation:  True
```

<!-- grade_id: cell-861772d45b70f49d, points: 2, grade: true, solution: false -->
```python
# DO NOT MODIFY, THIS IS UNIT TESTING!

assert string_rotations("abacd", "cdaba") == True
assert string_rotations("stackoverflow", "overflowstock") == False
assert string_rotations("0010100101101", "0010110100101") == True

# HIDDEN TEST
```

```python
# SAMPLE TEST, DO NOT MODIFY, JUST EXECUTE ONLY FOR OUTPUT
# You can add a cell below if you need to test a different input

print(string_rotations("abracad", "cadabra"))
```

Output:
```
True
```