# Programming Logic and Design Laboratory (LBYCPA1)

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or `YOUR ANSWER HERE`, as well as your name and other information below inside the double quotes. Remove the line `raise NotImplementedError()` once you have answered the problem.

REMINDERS: Section 4.15 of the Student Handbook (SH) 2018-2021 provides that:

As a Catholic University that prepares its students to become professionals and be the resource for Church and nation, De La Salle University puts a premium on academic honesty. A student's academic requirements, such as assignments, term papers, computer programs/projects, and thesis papers, should be his/her own work. He/She must distinguish between his/her own ideas and those of other authors.

The student must cite references, direct quotes, and other sources (including data obtained from tables, illustrations, figures, pictures, images, and video) following the prescribed format of the discipline. Should a student intend to submit a similar previous work for credit, in different courses, a written permission must first be obtained from the faculty members concerned. In addition, for works where he/she is only one of the authors, the following must be observed.

4.15.1 Provide a disclaimer stating that the work is not entirely his/hers;

4.15.2 The work cannot be published without written consent of the group of authors.

4.15.3 If the work is separable (i.e. parts are clearly attributed to the individual authors), the student may use and must cite his/her own part only; otherwise, the student must seek the written consent of the group of authors.

4.15.4 Computer codes obtained from open source reference may be copied freely, but the source should be cited, provided further that the faculty member allows the use of open source reference in the course.

## Rubric for the Final Laboratory Report

|Criteria|EXEMPLARY (90-100)|SATISFACTORY (80-89)|DEVELOPING (70-79)|BEGINNING (Below 70)|WEIGHT|
|:-|:-|:-|:-|:-|:-:|
|**Experimental Plan (Flowchart/ Algorithm)**|Experimental plan has supporting details and diagram/algorithm that is stated and well explained|Experimental plan has supporting details and diagram/algorithm that is stated but not explained|Experimental plan is vague or brief. It has supporting details and doesn’t have diagram/algorithm|No experimental plan presented|30%|
|**Conducting Experiment/ Activity**|Objective and Materials used (modern engineering tools, software, and instruments/equipment) are identified. Steps are easy to follow for conducting and experiment/activity|Objective is not stated. Materials used (modern engineering tools, software, and instruments/equipment) are identified. Steps are easy to follow for conducting and experiment/activity|Does not provide enough information to conduct an experiment/activity|No Objective, Materials used (modern engineering tools, software, and instruments/equipment), and steps for conducting experiment/activity provided|20%|
|**Codes/Data/ Program**|Data is well utilized in the program. Program codes are easy to read. Program output has no error. Questions are answered completely and correctly|Data is somewhat utilized in the program. Program code are easy to read. Program output has an output but logically incorrect. Some questions are answered completely and correctly|Data is not utilized in the program. It has a missing significant code/syntax in the program|No program presented|30%|
|**Grammar, logical presentation, and format**|The report was grammatically correct, logically presented and used the required format|The report had minimal grammatical errors and somewhat presented logically. The required format was used|The report had a lot of grammatical errors and not logically presented; the required format was barely used|The report had a lot of grammatical errors, was not logically presented and the required format was not used|20%|

```python
# Fill-up the following information (put inside double quotes)

NAME = ""                   # e.g. LIGUTAN, DINO DOMINIC F.
STUDENT_NUMBER = ""         # e.g. 20192118
SECTION = ""                # e.g. EQ1
DATE_TIME_COMPLETED = ""    # e.g. 03/21/2022,11:00
```

<hr style="border:1px solid black"> </hr>

# Problem Set 9
<hr style="border:1px solid gray"> </hr>

1. Write a function that reads any text file with a given filename and returns the count of a unique word contained in the text upon query. The function is case-sensitive (i.e., "The" is not the same as "the"). Use the `translate` and `maketrans` methods of strings to remove the punctuation marks of the read file. Given a the string `mystr`, the removal of punctuation marks can be done by the line:

        mystr = mystr.translate(str.maketrans('', '', string.punctuation))

    The function should return 0 if the query word is not found.

<!-- grade_id: cell-32216f35a45b784e, points: 0, grade: false, solution: true -->
```python
import string # import string for punctuation characters

def countWord(filename, query):
    if type(filename) is not str or type(query) is not str:
        raise TypeError("Both arguments must be a string")
        
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-a863d6410ce95310, points: 1, grade: true, solution: false -->
```python
# DO NOT MODIFY, THIS IS UNIT TESTING! [1 pt]
from nose.tools import assert_equal, assert_is_instance

assert_equal(countWord("alice29.txt", "she"), 497)
assert_equal(countWord("alice29.txt", "they"), 108)
assert_equal(countWord("alice29.txt", "Alice"), 385)
assert_is_instance(countWord("alice29.txt", ""), int)
```

<!-- grade_id: cell-b25b3cbf30198389, points: 1, grade: true, solution: false -->
```python
# HIDDEN TEST [1 pt]
```

2. Repeat the previous problem but instead of a query word, input a query number (a threshold) and return a Python dictionary that contains the words with its corresponding counts exceeding or equal to the query number.

<!-- grade_id: cell-2a67d9bb9ce95749, points: 0, grade: false, solution: true -->
```python
import string # import string for punctuation characters

def commonWords(filename, threshold):
    if type(filename) is not str:
        raise TypeError("First arguments must be a string")
    if type(threshold) is not int:
        raise TypeError("Second arguments must be an integer")
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-0ecdfae26dbb43fd, points: 1, grade: true, solution: false -->
```python
# DO NOT MODIFY, THIS IS UNIT TESTING! [1 pt]
from nose.tools import assert_equal, assert_is_instance

assert_equal(commonWords("alice29.txt", 500), {'the': 1514, 'to': 717, 'and': 774, 'a': 610})
assert_equal(commonWords("alice29.txt", 350), 
             {'I': 401, 'the': 1514, 'Alice': 385, 'was': 352, 'to': 717, 'of': 493, 'and': 774, 'she': 497, 'it': 482, 'in': 353, 'a': 610, 'said': 456})
assert_equal(commonWords("alice29.txt", 200), 
             {'I': 401, 'the': 1514, 'Alice': 385, 'was': 352, 'to': 717, 'of': 493, 'her': 243, 'and': 774, 'she': 497, 'it': 482, 'in': 353, 'a': 610, 'as': 246, 'that': 257, 'said': 456, 'you': 309})
assert_is_instance(commonWords("alice29.txt", 1000), dict)
```

<!-- grade_id: cell-3336c707423a9f21, points: 1, grade: true, solution: false -->
```python
# HIDDEN TEST [1 pt]
```

<!-- grade_id: cell-264fc6240e8bc614, points: 1, grade: true, solution: false -->
```python
# HIDDEN TEST [1 pt]
```

3. Create a function that will return a list of numbers from 1 up to a given number $n$ so that at all non-zero digits of the numbers divides itself. Provide exception handling for non-numerical values by returning the string "Specify a numerical input". (Note: You may need to read the built-in Python [Exceptions](https://docs.python.org/3/library/exceptions.html) to use the appropriate exception type)

<!-- grade_id: cell-f17ba2ccc79da39d, points: 0, grade: false, solution: true -->
```python
def divisibleByDigits(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-861772d45b70f49d, points: 1, grade: true, solution: false -->
```python
# DO NOT MODIFY, THIS IS UNIT TESTING! [1 pt]
from nose.tools import assert_equal, assert_is_instance

nums = divisibleByDigits(1000)
assert_equal(nums[15], 24)
assert_equal(nums[18], 36)
assert_equal(nums[25:35], [66, 70, 77, 80, 88, 90, 99, 100, 101, 102])
assert_equal(divisibleByDigits(30), [1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 15, 20, 22, 24, 30])
assert_equal(divisibleByDigits("Hello"), 'Specify a numerical input')
```

<!-- grade_id: cell-1fa0ead94c035a39, points: 1, grade: true, solution: false -->
```python
# HIDDEN TEST [1 pt]
```

4. The Intersection over Union (IoU) is computed as the intersection between two sets over their union. Determine the similarity between two programs with given filenames, read each of them into a word set (do not eliminate punctuation marks), and return the computed IoU value in percentage (round to two decimal places) as a string using the intersection and union sizes (e.g., `IoU("program1.cpp", "program2.cpp") = 100 # If exactly the same`). 

    Ensure exception handling if the files are inaccessible or cannot be opened as a text file. For files that cannot be accessed (`FileNotFoundError`), the function should return the string "Error: A file is not accessible". For files that cannot be opened as a text file (`UnicodeDecodeError`), the function should return the string "Error: A file cannot be opened as a text file". For all other exceptions, the function should not crash and return the corresponding exception message as a string. (Note: You may need to read the built-in Python [Exceptions](https://docs.python.org/3/library/exceptions.html) to use the appropriate exception type)

<!-- grade_id: cell-6ab996ab05760ce3, points: 0, grade: false, solution: true -->
```python
def IoU(filename1, filename2):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-52c031db984a5135, points: 1, grade: true, solution: false -->
```python
# DO NOT MODIFY, THIS IS UNIT TESTING! [1 pt]
from nose.tools import assert_equal, assert_is_instance

assert_equal(IoU("program1.c", "program2.c"), '35.43')
assert_equal(IoU("program2.c", "program1.c"), '35.43')
assert_equal(IoU("program1.c", "program1.c"), '100.00')
```

<!-- grade_id: cell-b88b738a393eb33d, points: 1, grade: true, solution: false -->
```python
# DO NOT MODIFY, THIS IS UNIT TESTING! [1 pt]
from nose.tools import assert_equal, assert_is_instance

assert_equal(IoU("program_1.c", "program2.c"), 'Error: A file is not accessible')
assert_equal(IoU("program1.c", "program_2.c"), 'Error: A file is not accessible')
assert_equal(IoU("program1.c", "python.c"), 'Error: A file cannot be opened as a text file')
```

<!-- grade_id: cell-b6381926685343b5, points: 1, grade: true, solution: false -->
```python
# HIDDEN TEST [1 pt]
```

***

<!-- grade_id: cell-4420f5b4060accd6, points: 0, grade: true, solution: false -->
```python

```