# Programming Logic and Design Laboratory (LBYCPA1)

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or `YOUR ANSWER HERE`, as well as your name and other information below inside the double quotes. Remove the line `raise NotImplementedError()` once you have answered the problem.

REMINDERS: Section 4.15 of the Student Handbook (SH) 2018-2021 provides that:

As a Catholic University that prepares its students to become professionals and be the resource for Church and nation, De La Salle University puts a premium on academic honesty. A student's academic requirements, such as assignments, term papers, computer programs/projects, and thesis papers, should be his/her own work. He/She must distinguish between his/her own ideas and those of other authors.

The student must cite references, direct quotes, and other sources (including data obtained from tables, illustrations, figures, pictures, images, and video) following the prescribed format of the discipline. Should a student intend to submit a similar previous work for credit, in different courses, a written permission must first be obtained from the faculty members concerned. In addition, for works where he/she is only one of the authors, the following must be observed.

4.15.1 Provide a disclaimer stating that the work is not entirely his/hers;

4.15.2 The work cannot be published without written consent of the group of authors.

4.15.3 If the work is separable (i.e. parts are clearly attributed to the individual authors), the student may use and must cite his/her own part only; otherwise, the student must seek the written consent of the group of authors.

4.15.4 Computer codes obtained from open source reference may be copied freely, but the source should be cited, provided further that the faculty member allows the use of open source reference in the course.

## Rubric for the Final Laboratory Report

|Criteria|EXEMPLARY (90-100)|SATISFACTORY (80-89)|DEVELOPING (70-79)|BEGINNING (Below 70)|WEIGHT|
|:-|:-|:-|:-|:-|:-:|
|**Experimental Plan (Flowchart/ Algorithm)**|Experimental plan has supporting details and diagram/algorithm that is stated and well explained|Experimental plan has supporting details and diagram/algorithm that is stated but not explained|Experimental plan is vague or brief. It has supporting details and doesn’t have diagram/algorithm|No experimental plan presented|30%|
|**Conducting Experiment/ Activity**|Objective and Materials used (modern engineering tools, software, and instruments/equipment) are identified. Steps are easy to follow for conducting and experiment/activity|Objective is not stated. Materials used (modern engineering tools, software, and instruments/equipment) are identified. Steps are easy to follow for conducting and experiment/activity|Does not provide enough information to conduct an experiment/activity|No Objective, Materials used (modern engineering tools, software, and instruments/equipment), and steps for conducting experiment/activity provided|20%|
|**Codes/Data/ Program**|Data is well utilized in the program. Program codes are easy to read. Program output has no error. Questions are answered completely and correctly|Data is somewhat utilized in the program. Program code are easy to read. Program output has an output but logically incorrect. Some questions are answered completely and correctly|Data is not utilized in the program. It has a missing significant code/syntax in the program|No program presented|30%|
|**Grammar, logical presentation, and format**|The report was grammatically correct, logically presented and used the required format|The report had minimal grammatical errors and somewhat presented logically. The required format was used|The report had a lot of grammatical errors and not logically presented; the required format was barely used|The report had a lot of grammatical errors, was not logically presented and the required format was not used|20%|

```python
# Fill-up the following information (put inside double quotes)

NAME = "DELA CRUZ, JOHN CARLO THEO S."                   # e.g. LIGUTAN, DINO DOMINIC F.
STUDENT_NUMBER = "12195758"         # e.g. 20192118
SECTION = "EQ1"                # e.g. EQ1
DATE_TIME_COMPLETED = "05/24/22, 10:43PM"    # e.g. 03/21/2022,11:00
```

<hr style="border:1px solid black"> </hr>

# Problem Set 7
<hr style="border:1px solid gray"> </hr>

1. Given a list of integers, write a function that will efficiently find the sum of the elements of a sublist that has the greatest value when all of its elements are summed together.

<!-- grade_id: cell-32216f35a45b784e, points: 0, grade: false, solution: true -->
```python
def maxsublistsum(nums):
    # This function finds the largest sum of the sublist possible from the given list. It should return an integer
    
    if type(nums) is not list:
        raise TypeError("Input must be a list")
        
    sum1 = 0
    maxSum = 0
    
    for x in nums:
        if sum1<0:
            sum1 = 0
        sum1 += x
        maxSum = max(maxSum,sum1)
    
    return int(maxSum)
```

<!-- grade_id: cell-a863d6410ce95310, points: 1, grade: true, solution: false -->
```python
# DO NOT MODIFY, THIS IS UNIT TESTING! [1 pt]
from nose.tools import assert_equal

assert_equal(maxsublistsum([1, 2, -5, 4, -3, 2]), 4)
assert_equal(maxsublistsum([1, 7, 8, -5, 3, 4, -3, 2, 7, -10, 1, 2, 1]), 24)
assert_equal(maxsublistsum([2, 3, -7, 8, 4, 5, -2, 3, 5]), 23)
```

<!-- grade_id: cell-b25b3cbf30198389, points: 1, grade: true, solution: false -->
```python
# HIDDEN TEST [1 pt]
```

```python
# SAMPLE TEST, DO NOT MODIFY! JUST EXECUTE ONLY FOR OUTPUT
# You can add a cell below if you need to test a different input

msg = input("Input a list of integers, separated by space: ")
nums = [int(num) for num in msg.split()]
print(f"Largest sum possible: {maxsublistsum(nums)}")
```

Output:
```
Input a list of integers, separated by space: 1 2 -5 4 -3 2
Largest sum possible: 4
```

2. Implement a function that will return a list of words that has number of letters greater than or equal to a certain given threshold and the number of such words as a tuple. It should accept an arbitrary string as its input. Punctuation marks should not be included for counting the length of a word.

<!-- grade_id: cell-2a67d9bb9ce95749, points: 0, grade: false, solution: true -->
```python
import string # import string for punctuation characters

def lists_words(msg, threshold=1):
    # This function returns a list of words whose length is greater than or equal to the threshold
    
    if type(msg) is not str:
        raise TypeError("Input not a valid string")
    
    msg = msg.translate(str.maketrans('', '', string.punctuation)) # removes all punctuations
    
    text = [x for x in msg.split() if len(x) >= threshold]
    return (text, len(text))

```

```python
zen_of_python = """
Beautiful is better than ugly.
Explicit is better than implicit.
Simple is better than complex.
Complex is better than complicated.
Flat is better than nested.
Sparse is better than dense.
Readability counts.
Special cases aren't special enough to break the rules.
Although practicality beats purity.
Errors should never pass silently.
Unless explicitly silenced.
In the face of ambiguity, refuse the temptation to guess.
There should be one-- and preferably only one --obvious way to do it.
Although that way may not be obvious at first unless you're Dutch.
Now is better than never.
Although never is often better than *right* now.
If the implementation is hard to explain, it's a bad idea.
If the implementation is easy to explain, it may be a good idea.
Namespaces are one honking great idea -- let's do more of those!
"""

lists_words(zen_of_python, 15)
```

Output:
```
([], 0)
```

<!-- grade_id: cell-0ecdfae26dbb43fd, points: 1, grade: true, solution: false -->
```python
# DO NOT MODIFY, THIS IS UNIT TESTING! [1 pt]
from nose.tools import assert_equal

zen_of_python = """
Beautiful is better than ugly.
Explicit is better than implicit.
Simple is better than complex.
Complex is better than complicated.
Flat is better than nested.
Sparse is better than dense.
Readability counts.
Special cases aren't special enough to break the rules.
Although practicality beats purity.
Errors should never pass silently.
Unless explicitly silenced.
In the face of ambiguity, refuse the temptation to guess.
There should be one-- and preferably only one --obvious way to do it.
Although that way may not be obvious at first unless you're Dutch.
Now is better than never.
Although never is often better than *right* now.
If the implementation is hard to explain, it's a bad idea.
If the implementation is easy to explain, it may be a good idea.
Namespaces are one honking great idea -- let's do more of those!
"""

assert_equal(lists_words(zen_of_python, 10), (['complicated','Readability','practicality','explicitly',
    'temptation','preferably','implementation', 'implementation','Namespaces'], 9))
assert_equal(lists_words(zen_of_python, 12), (['practicality', 'implementation', 'implementation'], 3))
assert_equal(lists_words(zen_of_python, 15), ([], 0))
```

<!-- grade_id: cell-3336c707423a9f21, points: 1, grade: true, solution: false -->
```python
# HIDDEN TEST [1 pt]
```

```python
# SAMPLE TEST, DO NOT MODIFY! JUST EXECUTE ONLY FOR OUTPUT
# You can add a cell below if you need to test a different input

threshold = int(input("Specify the minimum word length: "))
print(lists_words(zen_of_python, threshold))
```

Output:
```
Specify the minimum word length: 12
(['practicality', 'implementation', 'implementation'], 3)
```

3. An IP address is a set of four numbers that identifies a device on the internet or a local area network. These four numbers are separated by period (`.`) when written in decimal representation. The four numbers are whole numbers limited from 0 up to 255. An IP address is valid if it consists of four numbers (limited from 0 up to 255) with periods in between (e.g., `192.168.0.1`). Develop a function that will return a tuple of valid IP addresses in a given text.

<!-- grade_id: cell-f17ba2ccc79da39d, points: 0, grade: false, solution: true -->
```python
import string

txt1 = "Hello my ip address is 192.167.1.1 and I always listen to the radio 92.5. I know a lot about"\
        "IP addresses and that 19.255.244.212 is a valid ip and 256.300.12.12 is not."
txt2 = "Hackers get the IP address of people ... but I am not a hacker so I just ask them about it."\
        "Sometimes they say 2.3.4.1 or 19.23.110.45. Others say 90.4532.22.345 but I know its wrong."

def list_valid_ip(text):
    l1 = [i if i.count('.') == 3 and i[0].isdigit() else i.strip('.') if (i.count('.') == 4 and i[0].isdigit()) else '' for i in text.split()]
    l1 = [i for i in l1 if not i == '']
    l2 = [i.split('.') for i in l1]
    l3 = [l2[i] for i in range(len(l2)) if len(l2[i]) == 4]
    l4 = [l1[i] for i in range(len(l3)) if all(0 < int(y) < 256 for y in l3[i])]
    tupleList = tuple(l4)
    print(l1)
    print('\n')
    print(l2)
    print('\n')
    print(l3)
    print('\n')
    print(l4)
    print('\n')
    return tupleList
```

```python
sentence = 'Hello world a b c'
split_sentence = sentence.split(' ')
print(split_sentence)
```

Output:
```
['Hello', 'world', 'a', 'b', 'c']
```

<!-- grade_id: cell-861772d45b70f49d, points: 1, grade: true, solution: false -->
```python
# DO NOT MODIFY, THIS IS UNIT TESTING! [1 pt]
from nose.tools import assert_equal

assert_equal(list_valid_ip(txt1), ('192.167.1.1', '19.255.244.212'))
```

Output:
```
['192.167.1.1', '19.255.244.212', '256.300.12.12']


[['192', '167', '1', '1'], ['19', '255', '244', '212'], ['256', '300', '12', '12']]


[['192', '167', '1', '1'], ['19', '255', '244', '212'], ['256', '300', '12', '12']]


['192.167.1.1', '19.255.244.212']
```

<!-- grade_id: cell-1fa0ead94c035a39, points: 1, grade: true, solution: false -->
```python
# DO NOT MODIFY, THIS IS UNIT TESTING! [1 pt]
from nose.tools import assert_equal

assert_equal(list_valid_ip(txt2), ("2.3.4.1", "19.23.110.45"))
```

Output:
```
['2.3.4.1', '19.23.110.45', '90.4532.22.345']
[['2', '3', '4', '1'], ['19', '23', '110', '45'], ['90', '4532', '22', '345']]
[['2', '3', '4', '1'], ['19', '23', '110', '45'], ['90', '4532', '22', '345']]
['2.3.4.1', '19.23.110.45']
```

<!-- grade_id: cell-a79777fa5fed871f, points: 1, grade: true, solution: false -->
```python
# HIDDEN TEST [1 pt]
```

```python
# SAMPLE TEST, DO NOT MODIFY, JUST EXECUTE ONLY FOR OUTPUT
# You can add a cell below if you need to test a different input

print(list_valid_ip(txt1))
print(list_valid_ip(txt2))
```

Output:
```
['192.167.1.1', '19.255.244.212', '256.300.12.12']
[['192', '167', '1', '1'], ['19', '255', '244', '212'], ['256', '300', '12', '12']]
[['192', '167', '1', '1'], ['19', '255', '244', '212'], ['256', '300', '12', '12']]
['192.167.1.1', '19.255.244.212']
('192.167.1.1', '19.255.244.212')
['2.3.4.1', '19.23.110.45', '90.4532.22.345']
[['2', '3', '4', '1'], ['19', '23', '110', '45'], ['90', '4532', '22', '345']]
[['2', '3', '4', '1'], ['19', '23', '110', '45'], ['90', '4532', '22', '345']]
['2.3.4.1', '19.23.110.45']
('2.3.4.1', '19.23.110.45')
```

4. You need to implement two functions for this problem. For the first function: Given the list of student ID numbers and their section, implement a function that rearranges the list with each entry a tuple of the form `(ID, Section)` and then removes any duplicate entries. A tuple entry is a duplicate of another if both the ID and Section is the same for both tuples. For the second function: Given the list of tuples of the form `(ID, Section)` with no duplicates, implement a function that will output a list of tuples with a given section name. The tuple entries mut be sorted by student ID number.

<!-- grade_id: cell-36374e323cb18f1e, points: 0, grade: false, solution: true -->
```python
def reorganize_student_list(student_list):
    # This function reorganizes the list into tuples as well as eliminates any duplicate entries
    
    if type(student_list) is not list:
        raise TypeError("Input must be a list")
        
    #list to dictionary to tuples
    student_dct = {student_list[x]:student_list[x+1] for  x in range (0, len(student_list),2)}
    student_tup = [(idn, sec) for idn, sec in student_dct.items()]
    return student_tup
    

def sort_filter_student_list(student_list, section="EQ1"):
    # This function selects the tuples with the required section and returns a a list of tuples sorted by ID numbers
    
    if type(student_list) is not list:
        raise TypeError("Input must be a list")
    list1 = [x for x in student_list if x[1]==section]
    list1 = sorted(list1)
    
    return list1

```

```python
d1 = {120:"hi", 121:"nice"}
d1[119]="apex"
d1[120]="apex"
print(d1.items())
```

Output:
```
dict_items([(120, 'apex'), (121, 'nice'), (119, 'apex')])
```

```python
tup = [("a",1),("b",2)]
print(tup[1][1])
```

Output:
```
2
```

<!-- grade_id: cell-349cf941f53f1f52, points: 1, grade: true, solution: false -->
```python
# DO NOT MODIFY, THIS IS UNIT TESTING! [1 pt]
from nose.tools import assert_equal

studentList = [ 12143446, "EQ4", 12182311, "EQ1", 12198579, "EQ4", 12125946, "EQ1", 12195758, "EQ1",
    12111503, "EQ1", 12176451, "EQ7", 12124397, "EQ4", 12182311, "EQ1", 12174548, "EQ7",
    12120804, "EQ4", 12134775, "EQ4", 12117839, "EQ4", 12125601, "EQ1", 12184640, "EQ1",
    12184128, "EQ4", 12195766, "EQ1", 12176842, "EQ4", 12184128, "EQ4", 12111503, "EQ1",
    12117033, "EQ1", 12184985, "EQ7", 12125946, "EQ1", 12188867, "EQ7", 12143421, "EQ1",
    12182311, "EQ1", 12198323, "EQ7", 12112267, "EQ7", 12106844, "EQ1", 12112267, "EQ7",
    12143446, "EQ4", 12191043, "EQ1", 12171891, "EQ7", 12173371, "EQ1", 12191272, "EQ1",
    12174548, "EQ7", 12193143, "EQ7", 12190632, "EQ4", 12110760, "EQ4", 12198323, "EQ7",
    12198579, "EQ4", 12116602, "EQ4", 12183873, "EQ1", 12134759, "EQ7", 12133426, "EQ7",
    12110560, "EQ7", 12182621, "EQ1", 12116663, "EQ4", 12120804, "EQ4", 12198714, "EQ7",
    12143466, "EQ4", 12183873, "EQ1", 12107336, "EQ1", 12172146, "EQ1", 12111503, "EQ1",
    12110760, "EQ4", 12188867, "EQ7", 12107964, "EQ4", 12108880, "EQ4", 12194417, "EQ7",
    12143219, "EQ7", 12134759, "EQ7", 12179310, "EQ4", 12184128, "EQ4", 12142158, "EQ7", 
    12198579, "EQ4", 12194417, "EQ7", 12176842, "EQ4", 12179620, "EQ4", 12109080, "EQ1"]

assert_equal(reorganize_student_list(studentList), [(12143446, 'EQ4'), (12182311, 'EQ1'), (12198579, 'EQ4'), (12125946, 'EQ1'),
    (12195758, 'EQ1'), (12111503, 'EQ1'), (12176451, 'EQ7'), (12124397, 'EQ4'), (12174548, 'EQ7'), (12120804, 'EQ4'),
    (12134775, 'EQ4'), (12117839, 'EQ4'), (12125601, 'EQ1'), (12184640, 'EQ1'), (12184128, 'EQ4'), (12195766, 'EQ1'),
    (12176842, 'EQ4'), (12117033, 'EQ1'), (12184985, 'EQ7'), (12188867, 'EQ7'), (12143421, 'EQ1'), (12198323, 'EQ7'),
    (12112267, 'EQ7'), (12106844, 'EQ1'), (12191043, 'EQ1'), (12171891, 'EQ7'), (12173371, 'EQ1'), (12191272, 'EQ1'),
    (12193143, 'EQ7'), (12190632, 'EQ4'), (12110760, 'EQ4'), (12116602, 'EQ4'), (12183873, 'EQ1'), (12134759, 'EQ7'),
    (12133426, 'EQ7'), (12110560, 'EQ7'), (12182621, 'EQ1'), (12116663, 'EQ4'), (12198714, 'EQ7'), (12143466, 'EQ4'),
    (12107336, 'EQ1'), (12172146, 'EQ1'), (12107964, 'EQ4'), (12108880, 'EQ4'), (12194417, 'EQ7'), (12143219, 'EQ7'),
    (12179310, 'EQ4'), (12142158, 'EQ7'), (12179620, 'EQ4'), (12109080, 'EQ1')])
assert_equal(len(reorganize_student_list(studentList)), 50)
```

<!-- grade_id: cell-bb9ae99ac511258f, points: 1, grade: true, solution: false -->
```python
# DO NOT MODIFY, THIS IS UNIT TESTING! [1 pt]
from nose.tools import assert_equal

organized_student_list = reorganize_student_list(studentList)
assert_equal(sort_filter_student_list(organized_student_list, section="EQ1"),
    [(12106844, 'EQ1'), (12107336, 'EQ1'), (12109080, 'EQ1'), (12111503, 'EQ1'), (12117033, 'EQ1'), (12125601, 'EQ1'),
     (12125946, 'EQ1'), (12143421, 'EQ1'), (12172146, 'EQ1'), (12173371, 'EQ1'), (12182311, 'EQ1'), (12182621, 'EQ1'),
     (12183873, 'EQ1'), (12184640, 'EQ1'), (12191043, 'EQ1'), (12191272, 'EQ1'), (12195758, 'EQ1'), (12195766, 'EQ1')])
assert_equal(sort_filter_student_list(organized_student_list, section="EQ4"),
    [(12107964, 'EQ4'), (12108880, 'EQ4'), (12110760, 'EQ4'), (12116602, 'EQ4'), (12116663, 'EQ4'), (12117839, 'EQ4'),
     (12120804, 'EQ4'), (12124397, 'EQ4'), (12134775, 'EQ4'), (12143446, 'EQ4'), (12143466, 'EQ4'), (12176842, 'EQ4'),
     (12179310, 'EQ4'), (12179620, 'EQ4'), (12184128, 'EQ4'), (12190632, 'EQ4'), (12198579, 'EQ4')])
assert_equal(sort_filter_student_list(organized_student_list, section="EQ7"),
    [(12110560, 'EQ7'), (12112267, 'EQ7'), (12133426, 'EQ7'), (12134759, 'EQ7'), (12142158, 'EQ7'), (12143219, 'EQ7'),
     (12171891, 'EQ7'), (12174548, 'EQ7'), (12176451, 'EQ7'), (12184985, 'EQ7'), (12188867, 'EQ7'), (12193143, 'EQ7'),
     (12194417, 'EQ7'), (12198323, 'EQ7'), (12198714, 'EQ7')])
assert_equal(sort_filter_student_list(organized_student_list, section="EQ8"), [])
```

<!-- grade_id: cell-b88b738a393eb33d, points: 1, grade: true, solution: false -->
```python
# HIDDEN TEST [1 pt]
```

```python
# SAMPLE TEST, DO NOT MODIFY, JUST EXECUTE ONLY FOR OUTPUT
# You can add a cell below if you need to test a different input

print(reorganize_student_list(studentList))
print(sort_filter_student_list(reorganize_student_list(studentList), section="EQ1"))
print(sort_filter_student_list(reorganize_student_list(studentList), section="EQ4"))
print(sort_filter_student_list(reorganize_student_list(studentList), section="EQ7"))
```

Output:
```
[(12143446, 'EQ4'), (12182311, 'EQ1'), (12198579, 'EQ4'), (12125946, 'EQ1'), (12195758, 'EQ1'), (12111503, 'EQ1'), (12176451, 'EQ7'), (12124397, 'EQ4'), (12174548, 'EQ7'), (12120804, 'EQ4'), (12134775, 'EQ4'), (12117839, 'EQ4'), (12125601, 'EQ1'), (12184640, 'EQ1'), (12184128, 'EQ4'), (12195766, 'EQ1'), (12176842, 'EQ4'), (12117033, 'EQ1'), (12184985, 'EQ7'), (12188867, 'EQ7'), (12143421, 'EQ1'), (12198323, 'EQ7'), (12112267, 'EQ7'), (12106844, 'EQ1'), (12191043, 'EQ1'), (12171891, 'EQ7'), (12173371, 'EQ1'), (12191272, 'EQ1'), (12193143, 'EQ7'), (12190632, 'EQ4'), (12110760, 'EQ4'), (12116602, 'EQ4'), (12183873, 'EQ1'), (12134759, 'EQ7'), (12133426, 'EQ7'), (12110560, 'EQ7'), (12182621, 'EQ1'), (12116663, 'EQ4'), (12198714, 'EQ7'), (12143466, 'EQ4'), (12107336, 'EQ1'), (12172146, 'EQ1'), (12107964, 'EQ4'), (12108880, 'EQ4'), (12194417, 'EQ7'), (12143219, 'EQ7'), (12179310, 'EQ4'), (12142158, 'EQ7'), (12179620, 'EQ4'), (12109080, 'EQ1')]
[(12106844, 'EQ1'), (12107336, 'EQ1'), (12109080,
```

***

<!-- grade_id: cell-4420f5b4060accd6, points: 0, grade: true, solution: false -->
```python

```