# Programming Logic and Design Laboratory (LBYCPA1)

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or `YOUR ANSWER HERE`, as well as your name and other information below inside the double quotes. Remove the line `raise NotImplementedError()` once you have answered the problem.

REMINDERS: Section 4.15 of the Student Handbook (SH) 2018-2021 provides that:

As a Catholic University that prepares its students to become professionals and be the resource for Church and nation, De La Salle University puts a premium on academic honesty. A student's academic requirements, such as assignments, term papers, computer programs/projects, and thesis papers, should be his/her own work. He/She must distinguish between his/her own ideas and those of other authors.

The student must cite references, direct quotes, and other sources (including data obtained from tables, illustrations, figures, pictures, images, and video) following the prescribed format of the discipline. Should a student intend to submit a similar previous work for credit, in different courses, a written permission must first be obtained from the faculty members concerned. In addition, for works where he/she is only one of the authors, the following must be observed.

4.15.1 Provide a disclaimer stating that the work is not entirely his/hers;

4.15.2 The work cannot be published without written consent of the group of authors.

4.15.3 If the work is separable (i.e. parts are clearly attributed to the individual authors), the student may use and must cite his/her own part only; otherwise, the student must seek the written consent of the group of authors.

4.15.4 Computer codes obtained from open source reference may be copied freely, but the source should be cited, provided further that the faculty member allows the use of open source reference in the course.

## Rubric for the Final Laboratory Report

|Criteria|EXEMPLARY (90-100)|SATISFACTORY (80-89)|DEVELOPING (70-79)|BEGINNING (Below 70)|WEIGHT|
|:-|:-|:-|:-|:-|:-:|
|**Experimental Plan (Flowchart/ Algorithm)**|Experimental plan has supporting details and diagram/algorithm that is stated and well explained|Experimental plan has supporting details and diagram/algorithm that is stated but not explained|Experimental plan is vague or brief. It has supporting details and doesn’t have diagram/algorithm|No experimental plan presented|30%|
|**Conducting Experiment/ Activity**|Objective and Materials used (modern engineering tools, software, and instruments/equipment) are identified. Steps are easy to follow for conducting and experiment/activity|Objective is not stated. Materials used (modern engineering tools, software, and instruments/equipment) are identified. Steps are easy to follow for conducting and experiment/activity|Does not provide enough information to conduct an experiment/activity|No Objective, Materials used (modern engineering tools, software, and instruments/equipment), and steps for conducting experiment/activity provided|20%|
|**Codes/Data/ Program**|Data is well utilized in the program. Program codes are easy to read. Program output has no error. Questions are answered completely and correctly|Data is somewhat utilized in the program. Program code are easy to read. Program output has an output but logically incorrect. Some questions are answered completely and correctly|Data is not utilized in the program. It has a missing significant code/syntax in the program|No program presented|30%|
|**Grammar, logical presentation, and format**|The report was grammatically correct, logically presented and used the required format|The report had minimal grammatical errors and somewhat presented logically. The required format was used|The report had a lot of grammatical errors and not logically presented; the required format was barely used|The report had a lot of grammatical errors, was not logically presented and the required format was not used|20%|

```python
# Fill-up the following information (put inside double quotes)

NAME = "DELA CRUZ, JOHN CARLO THEO S."                   # e.g. LIGUTAN, DINO DOMINIC F.
STUDENT_NUMBER = "12195758"         # e.g. 20192118
SECTION = "EQ1"                # e.g. EQ1
DATE_TIME_COMPLETED = "05/06/2022, 12:05AM"    # e.g. 03/21/2022,11:00
```

<hr style="border:1px solid black"> </hr>

# Problem Set 6
<hr style="border:1px solid gray"> </hr>

1. An message encryption method which appears similar to a Ceasar cipher substitution is to be implemented in the following way: each of the characters of the alphabet is re-mapped such that the ordering of the alphabet is reversed. This means that "A" becomes a "Z", "B" becomes a "Y", "W" becomes a "C" and so on. The same re-mapping is done for the numbers such that a "0" becomes a "9", an "8" becomes a "2" and so on. Implement a function that takes a sentence as a string and returns the encrypted message as a string.

<!-- grade_id: cell-32216f35a45b784e, points: 0, grade: false, solution: true -->
```python
def encryptMessage(message):
    # This function encrypts a message such that each character is substituted with its equivalent reversed alphabet ordering.

    if type(message) is not str:
        raise TypeError("Input a string only")
    
    encrypt = str()
    for i in message:
        if i.isalpha():
            num = ord(i)
            if i.isupper():
                num = (ord('Z')-num) + ord('A')
                encrypt += chr(num)
            elif i.islower:
                num = (ord('z')-num) + ord('a')
                encrypt += chr(num)
        elif i.isdecimal():
            num = ord(i)
            num = (ord('9')-num) + ord ('0')
            encrypt += chr(num)
        else:
            num = ord(i)
            encrypt += chr(num)
    return encrypt
        
```

<!-- grade_id: cell-a863d6410ce95310, points: 1, grade: true, solution: false -->
```python
# DO NOT MODIFY, THIS IS UNIT TESTING! [1 pt]
from nose.tools import assert_equal

assert_equal(encryptMessage('ABCDEFGHIJKLMNOPQRSTUVWXYZ'), 'ZYXWVUTSRQPONMLKJIHGFEDCBA')
assert_equal(encryptMessage('abcdefghijklmnopqrstuvwxyz'), 'zyxwvutsrqponmlkjihgfedcba')
assert_equal(encryptMessage('0123456789'),'9876543210')
assert_equal(encryptMessage('Hello, World!'), 'Svool, Dliow!')
assert_equal(encryptMessage('Let us see if you can still read this!'), 'Ovg fh hvv ru blf xzm hgroo ivzw gsrh!')
```

<!-- grade_id: cell-b25b3cbf30198389, points: 1, grade: true, solution: false -->
```python
# HIDDEN TEST [1 pt]
```

```python
# SAMPLE TEST, DO NOT MODIFY! JUST EXECUTE ONLY FOR OUTPUT
# You can add a cell below if you need to test a different input

msg = input("Type your message: ")
print(f'Your encrypted message: "{encryptMessage(msg)}"')
```

Output:
```
Type your message: Hotdog
Your encrypted message: "Slgwlt"
```

2. The bitwise logical operation **XNOR** (exclusive-NOR) takes two strings and compares each character in their respective positions as follows: if exactly one of the characters have a "1", the corresponding character in the resulting string is a "0"; otherwise it is a "1". Example: "10010001" XNOR "10010010" becomes "11111100". Implement a Python function that accepts arbitrary string length composed of "0"s and "1"s only. Both string inputs must have the same length. For reference, the truth table for the XNOR operation is as follows:
| A | B | XNOR(A, B) |
|:-:|:-:|:-:|
| 0 | 0 | 1 |
| 0 | 1 | 0 |
| 1 | 0 | 0 |
| 1 | 1 | 1 |

<!-- grade_id: cell-2a67d9bb9ce95749, points: 0, grade: false, solution: true -->
```python
def bitwise_XNOR(string1, string2):
    # This function performs the logical bitwise XNOR on two n-character strings
    
    checkInput(string1)
    checkInput(string2)

    if len(string1) != len(string2):
        raise ValueError("String inputs must have same lengths")
    result = str()
    for char1, char2 in zip(string1, string2):
        if char1 == '0' and char2 == '0': 
            result += '1'
        elif char1 == '1' and char2 =='1':
            result +='1'
        else:
            result +='0'
            
    return result
def checkInput(string):
    if type(string) is not str:
        raise TypeError("Input not a valid string")
    for char in string:
        if not (char == "0" or char == "1"):
            raise ValueError("Strings must contain a '0' or '1' characters only")
```

<!-- grade_id: cell-0ecdfae26dbb43fd, points: 1, grade: true, solution: false -->
```python
# DO NOT MODIFY, THIS IS UNIT TESTING! [1 pt]
from nose.tools import assert_equal

assert_equal(bitwise_XNOR("0010011011", "1010000011"), "0111100111")
assert_equal(bitwise_XNOR("11100000", "00011111"), "00000000")
assert_equal(bitwise_XNOR("1010101", "1111000"), "1010010")
assert_equal(bitwise_XNOR("1", "0"), "0")
```

<!-- grade_id: cell-3336c707423a9f21, points: 1, grade: true, solution: false -->
```python
# HIDDEN TEST [1 pt]
```

```python
# SAMPLE TEST, DO NOT MODIFY! JUST EXECUTE ONLY FOR OUTPUT
# You can add a cell below if you need to test a different input

bitwise_XNOR("0110110010010010", "1111110000110111")
```

Output:
```
'0110111101011010'
```

3. A border made up of a specified character is to be printed with specified side length. Implement a function that accepts two arguments: number of characters per side and the character to be used as a border. The function must return a **string**. Each line must end with the newline (`\n`) character.

<!-- grade_id: cell-5bf885b89699825d, points: 0, grade: false, solution: true -->
```python
def printBorder(width, char='*'):
    # Prints a border of specified width using a single character
    
    if type(char) is not str:
        raise TypeError("Input character must be a string")
    if type(width) is not int:
        raise TypeError("Input width must be an integer")
    if len(char) < 1 or len(char) > 1:
        raise ValueError("Specify single character input")
    if width < 0:
        raise ValueError("Specify a non-negative integer input")
        
    border = str()
    for char1 in char:
        for i in range(width):
            for j in range(width):
                if i == 0 or i == width - 1:
                    border += char1
                else:
                    if j == 0 or j == width - 1:
                        border += char1
                    else: 
                        border +=' '
            border +='\n'
                        
    return border
```

<!-- grade_id: cell-67877fb08827f061, points: 1, grade: true, solution: false -->
```python
# DO NOT MODIFY, THIS IS UNIT TESTING! [1 pt]
from nose.tools import assert_equal

assert_equal(printBorder(3), \
'***\n'\
'* *\n'\
'***\n')

assert_equal(printBorder(5, char='#'), \
'#####\n'\
'#   #\n'\
'#   #\n'\
'#   #\n'\
'#####\n')

assert_equal(printBorder(7, char='&'), \
'&&&&&&&\n'\
'&     &\n'\
'&     &\n'\
'&     &\n'\
'&     &\n'\
'&     &\n'\
'&&&&&&&\n')
```

<!-- grade_id: cell-54489cd2fa804445, points: 1, grade: true, solution: false -->
```python
# HIDDEN TEST [1 pt]
```

```python
# SAMPLE TEST, DO NOT MODIFY! JUST EXECUTE ONLY FOR OUTPUT

size = int(input("What border size? "))
print(printBorder(size))
```

Output:
```
What border size? 3
***
* *
***
```

4. Write a function that will emulate the DLSU email address generating system given that the input is a list of names in a string separated by commas. The output must be a string of comma-separated (`, `) email addresses (with space in between also) similar to your DLSU email address. (Refer to the unit testing cell for the expected output format)

<!-- grade_id: cell-ddf77fa3ab60f546, points: 0, grade: false, solution: true -->
```python
def namesToEmails(names):
    # Emulates the DLSU email address generation system and returns a list of email addresses as strings
    
    names = names.lower() 
    email = "{}_{}@dlsu.edu.ph" 
    nameslist = names.split(',') 
    namesspace = [x.strip() for x in nameslist] 
    nameslist2 = [x.split(' ') for x in namesspace] 
    namestoEmails = [email.format (x[0], x[1]) for x in nameslist2] 
    #for x in nameslist2: 
    #    namestoEmails.append(email.format(x[0], x[1]))
    
    result = ', '.join(namestoEmails) #join list 
    
    return result
    
#     print("nameslist: ", nameslist ,"\n")
#     print("namesspace: ", namesspace,"\n")
#     print("nameslist2: ", nameslist2,"\n")
#     print("namesEmails: ", namestoEmails,"\n")
```

<!-- grade_id: cell-9fb1a74bdba298aa, points: 1, grade: true, solution: false -->
```python
# DO NOT MODIFY, THIS IS UNIT TESTING! [1 pt]
from nose.tools import assert_equal

names = "Lauren Bassett, Dante Green, Rida Mohamed, William Livingston, Cora Callahan, "\
"Deen Manning, Taiba Chapman, Eiliyah Boyle, Brad Mullen, Lyndon Devine"

assert_equal(namesToEmails(names), "lauren_bassett@dlsu.edu.ph, dante_green@dlsu.edu.ph, "\
"rida_mohamed@dlsu.edu.ph, william_livingston@dlsu.edu.ph, "\
"cora_callahan@dlsu.edu.ph, deen_manning@dlsu.edu.ph, "\
"taiba_chapman@dlsu.edu.ph, eiliyah_boyle@dlsu.edu.ph, "\
"brad_mullen@dlsu.edu.ph, lyndon_devine@dlsu.edu.ph")
```

<!-- grade_id: cell-bc79be717a9ac2ab, points: 1, grade: true, solution: false -->
```python
# HIDDEN TEST [1 pt]
```

```python
# SAMPLE TEST, DO NOT MODIFY, JUST EXECUTE ONLY FOR OUTPUT
# You can add a cell below if you need to test a different input

print(namesToEmails("Sumaya Paine, Melinda Colley, Rafe Vang, Edan Mcdonald, Simra Robles, "\
      "Abul Johnson, Hajrah Kaur, Bryan Davidson, Clarence Hogg, Ajwa Sears"))
```

Output:
```
sumaya_paine@dlsu.edu.ph, melinda_colley@dlsu.edu.ph, rafe_vang@dlsu.edu.ph, edan_mcdonald@dlsu.edu.ph, simra_robles@dlsu.edu.ph, abul_johnson@dlsu.edu.ph, hajrah_kaur@dlsu.edu.ph, bryan_davidson@dlsu.edu.ph, clarence_hogg@dlsu.edu.ph, ajwa_sears@dlsu.edu.ph
```

5. Given a list of students and their scores for three exams as comma-separated values, develop a function that will calculate the average score of the exams of each student. First, complete the function `compute_average()` that will return the average score rounded to 1 decimal place as a float by accepting a string of comma-separated grades. Then, complete the function `combine_averages()` that will convert a student entry into the following format: `Name: Average score\n`. Note that each entry should have a newline character (`\n`) at the end and the functiom must return a single string.

<!-- grade_id: cell-f17ba2ccc79da39d, points: 0, grade: false, solution: true -->
```python
# DO NOT MODIFY, THIS IS UNIT TESTING! [1 pt]
student_scores = \
"""Anna,60,68,89
Darcey,80,78,90
Dylan,90,78,95
Eva,78,76,67
Layla,89,88,95
Lewis,78,88,89
Lucas,50,65,45
Theodore,88,86,87
Victoria,55,65,45
Zoe,94,67,96
"""
# "60, 68, 69"
# average
def compute_average(grade_string):
    # This function accepts a string with three comma-separated grades and returns the average as a float
    
    if type(grade_string) is not str:
        raise TypeError("Input must be a string")
    grade_string = grade_string.strip(',') # tanggal yung specific chuchu
    grade_string = grade_string.split(',') # string to array
    integers = [int(x) for x in grade_string]
    average = round(((integers[0] + integers [1] + integers[2]) / 3),1)
    return average 

# student_scores = string
# name, compute_average
def combine_averages(student_table):
    # This function combines all student average scores into a string. 
    # It should call the compute_average() function to calculate the average
    
    if type(student_table) is not str:
        raise TypeError("Input must be a string")
        
    result = "{}: {}"
    student_table = student_table.strip('\n') 
    student_table = student_table.split('\n')    
    student_table_space = [x.strip(',') for x in student_table]
    combine = [x.split(',') for x in student_table_space]
    # avgList = [x.split(',') for x in combine]
    List1 = [compute_average(','.join([x[1],x[2],x[3]]))for x in combine]
    List2 = [x[0] for x in combine]
    averageTable = [result.format (List2[x],List1[x]) for x in range (0,len(combine))] 
    
    # for x in range (0,len(combine))
    #     result.format(List2[x],List1[x])
    
    end ='\n'.join(averageTable)
    
    return end + '\n'

```

```python
#test
student_table = """Anna,60,68,89
Darcey,80,78,90
Dylan,90,78,95
Eva,78,76,67
Layla,89,88,95
Lewis,78,88,89
Lucas,50,65,45
Theodore,88,86,87
Victoria,55,65,45
Zoe,94,67,96
"""
result = "{}: {}"
student_table = student_table.strip('\n') 
student_table = student_table.split('\n')    
student_table_space = [x.strip(',') for x in student_table]
combine = [x.split(',') for x in student_table_space]
# avgList = [x.split(',') for x in combine]
List1 = [compute_average(','.join([x[1],x[2],x[3]]))for x in combine]
List2 = [x[0] for x in combine]
averageTable = [result.format (List2[x],List1[x]) for x in range (0,len(combine))] 
# for x in range (0,len(combine))
#     result.format(List2[x],List1[x])
end ='\n'.join(averageTable)

print(student_table)
print("\n")
print(student_table_space)
print("\n")
print(combine)
print("\n")
print(end)


```

Output:
```
['Anna,60,68,89', 'Darcey,80,78,90', 'Dylan,90,78,95', 'Eva,78,76,67', 'Layla,89,88,95', 'Lewis,78,88,89', 'Lucas,50,65,45', 'Theodore,88,86,87', 'Victoria,55,65,45', 'Zoe,94,67,96']


['Anna,60,68,89', 'Darcey,80,78,90', 'Dylan,90,78,95', 'Eva,78,76,67', 'Layla,89,88,95', 'Lewis,78,88,89', 'Lucas,50,65,45', 'Theodore,88,86,87', 'Victoria,55,65,45', 'Zoe,94,67,96']


[['Anna', '60', '68', '89'], ['Darcey', '80', '78', '90'], ['Dylan', '90', '78', '95'], ['Eva', '78', '76', '67'], ['Layla', '89', '88', '95'], ['Lewis', '78', '88', '89'], ['Lucas', '50', '65', '45'], ['Theodore', '88', '86', '87'], ['Victoria', '55', '65', '45'], ['Zoe', '94', '67', '96']]


Anna: 72.3
Darcey: 82.7
Dylan: 87.7
Eva: 73.7
Layla: 90.7
Lewis: 85.0
Lucas: 53.3
Theodore: 87.0
Victoria: 55.0
Zoe: 85.7
```

```python
name = "art, hotdog, komz, bleasion, opiad"

print(name.split())
print(name.strip())
```

Output:
```
['art,', 'hotdog,', 'komz,', 'bleasion,', 'opiad']
art, hotdog, komz, bleasion, opiad
```

<!-- grade_id: cell-861772d45b70f49d, points: 1, grade: true, solution: false -->
```python
# DO NOT MODIFY, THIS IS UNIT TESTING! [1 pt]
from nose.tools import assert_equal

assert_equal(compute_average(",60,68,89"), 72.3)
assert_equal(compute_average("90,78,95"), 87.7)
assert_equal(compute_average(",78,76,67,"), 73.7)

# HIDDEN TEST
```

<!-- grade_id: cell-a79777fa5fed871f, points: 1, grade: true, solution: false -->
```python
# DO NOT MODIFY, THIS IS UNIT TESTING! [1 pt]
from nose.tools import assert_equal

assert_equal(combine_averages(student_scores), \
"""Anna: 72.3
Darcey: 82.7
Dylan: 87.7
Eva: 73.7
Layla: 90.7
Lewis: 85.0
Lucas: 53.3
Theodore: 87.0
Victoria: 55.0
Zoe: 85.7
""")

# HIDDEN TEST
```

```python
# SAMPLE TEST, DO NOT MODIFY, JUST EXECUTE ONLY FOR OUTPUT
# You can add a cell below if you need to test a different input

print(combine_averages(student_scores))
```

Output:
```
Anna: 72.3
Darcey: 82.7
Dylan: 87.7
Eva: 73.7
Layla: 90.7
Lewis: 85.0
Lucas: 53.3
Theodore: 87.0
Victoria: 55.0
Zoe: 85.7
```

***

<!-- grade_id: cell-4420f5b4060accd6, points: 0, grade: true, solution: false -->
```python

```